// Numbas version: exam_results_page_options
{"name":"Vraag voorbeeld","duration":0,"showQuestionGroupNames":false,"showstudentname":true,"allowPrinting":true,"resources":[],"extensions":[],"custom_part_types":[],"navigation":{"startpassword":"","allowregen":true,"navigatemode":"sequence","onleave":{"action":"none","message":""},"showresultspage":"oncompletion","reverse":true,"browse":true,"allowsteps":true,"showfrontpage":false,"preventleave":true},"timing":{"allowPause":true,"timeout":{"action":"none","message":""},"timedwarning":{"action":"none","message":""}},"feedback":{"showactualmark":true,"showtotalmark":true,"showanswerstate":true,"allowrevealanswer":true,"reviewshowscore":true,"reviewshowfeedback":true,"reviewshowexpectedanswer":true,"reviewshowadvice":true,"advice":"","intro":"","feedbackmessages":[]},"question_groups":[{"name":"Q1","pickingStrategy":"random-subset","pickQuestions":1,"questions":[{"name":"percentages-shoes-without_templates_en","statement":"You want to buy a pair of shoes, but wait for the sales to start.\n<br/>\nOriginally, a pair of shoes was priced $\\var{price_orig}$ euro.\nYou now get a $\\var{percent}$% reduction.","advice":"","parts":[{"type":"gapfill","prompt":"How much are the shoes now? [[0]]","marks":0,"steps":[],"stepsPenalty":0,"showCorrectAnswer":true,"showFeedbackIcon":true,"useCustomName":true,"customName":"","variableReplacements":[],"variableReplacementStrategy":"originalfirst","adaptiveMarkingPenalty":0,"customMarkingAlgorithm":"","extendBaseMarkingAlgorithm":true,"sortAnswers":false,"gaps":[{"type":"numberentry","prompt":"","marks":1,"steps":[],"stepsPenalty":0,"showCorrectAnswer":true,"showFeedbackIcon":true,"useCustomName":true,"customName":"","variableReplacements":[],"variableReplacementStrategy":"originalfirst","adaptiveMarkingPenalty":0,"customMarkingAlgorithm":"","extendBaseMarkingAlgorithm":true,"correctAnswerFraction":true,"correctAnswerStyle":"plain","allowFractions":true,"notationStyles":["plain"],"mustBeReduced":false,"mustBeReducedPC":0,"precisionType":"none","precision":0,"precisionPartialCredit":0,"precisionMessage":"","strictPrecision":true,"showPrecisionHint":true,"showFractionHint":true,"answer":"price_now"}]}],"builtin_constants":{"e":true,"i":true,"pi,π":true},"constants":[],"variables":{"percent":{"name":"percent","definition":"random(20,30,40,50,60)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"price_now":{"name":"price_now","definition":"price_orig*(1-percent/100)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"price_orig":{"name":"price_orig","definition":"random(15,20,25,30,35)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false}},"variablesTest":{"condition":"","maxRuns":100},"functions":{},"ungrouped_variables":["percent","price_now","price_orig"],"variable_groups":[],"rulesets":{"std":["basic","unitFactor","unitPower","unitDenominator","zeroFactor","zeroTerm","zeroPower","zeroBase","constantsFirst","sqrtProduct","sqrtDivision","sqrtSquare","otherNumbers","simplifyFractions","trig","cancelTerms","cancelFactors","collectLikeFractions"]},"preamble":{"js":"","css":"tt { font-family:monospace !important; background:lightgray !important}\n.tt { font-family:monospace !important; background:lightgray !important}\nul.info { text-indent: -28px; margin-left: 28px;}\nh4 { display: inline;}    /* prevent newline after a), b), ... in parts; */\ntable { margin-left: unset; }     /* tables flush-left ... */\ntable.open td { padding-bottom: 4rem !important; }  /* bigger padding in template-OPENx */\nli.instruction {\n  /* background-color: lightgray; */\n  /* list-style-type: none; */\n  /* display: inline-block;   /* reset list-item */\n  font-size: smaller;\n  margin-left: 2rem;\n  color: gray;\n  list-style-position: outside;\n  /* border: gray solid; */\n  /* padding: 0.5rem;    */\n}\n"},"navigation":{"allowregen":true,"showfrontpage":false,"preventleave":false},"extensions":[],"resources":[],"tags":[],"custom_part_types":[]}]}]}