// Numbas version: exam_results_page_options
{"name":"Templatability","duration":0,"showQuestionGroupNames":false,"showstudentname":true,"allowPrinting":true,"resources":[],"extensions":[],"custom_part_types":[],"navigation":{"startpassword":"","allowregen":true,"navigatemode":"sequence","onleave":{"action":"none","message":""},"showresultspage":"oncompletion","reverse":true,"browse":true,"allowsteps":true,"showfrontpage":false,"preventleave":true},"timing":{"allowPause":true,"timeout":{"action":"none","message":""},"timedwarning":{"action":"none","message":""}},"feedback":{"showactualmark":true,"showtotalmark":true,"showanswerstate":true,"allowrevealanswer":true,"reviewshowscore":true,"reviewshowfeedback":true,"reviewshowexpectedanswer":true,"reviewshowadvice":true,"advice":"","intro":"","feedbackmessages":[]},"question_groups":[{"name":"Vraagstukken","pickingStrategy":"all-ordered","questions":[{"name":"percentages-shoes.yaml","statement":"You want to buy a pair of shoes, but wait for the sales to start.\n<li class=\"instruction\">De vraag zoals gedefinieerd in questions</li>","advice":"{advices}","parts":[{"type":"gapfill","prompt":"{prompts[0]}  [[0]]","marks":0,"steps":[],"stepsPenalty":0,"showCorrectAnswer":true,"showFeedbackIcon":true,"useCustomName":true,"customName":"","variableReplacements":[],"variableReplacementStrategy":"originalfirst","adaptiveMarkingPenalty":0,"customMarkingAlgorithm":"","extendBaseMarkingAlgorithm":true,"sortAnswers":false,"gaps":[{"type":"numberentry","prompt":"","marks":1,"steps":[],"stepsPenalty":0,"showCorrectAnswer":true,"showFeedbackIcon":true,"useCustomName":true,"customName":"","variableReplacements":[],"variableReplacementStrategy":"originalfirst","adaptiveMarkingPenalty":0,"customMarkingAlgorithm":"","extendBaseMarkingAlgorithm":true,"correctAnswerFraction":false,"correctAnswerStyle":"plain","allowFractions":true,"notationStyles":["plain"],"mustBeReduced":false,"mustBeReducedPC":0,"precisionType":"none","precision":0,"precisionPartialCredit":0,"precisionMessage":"","strictPrecision":true,"showPrecisionHint":true,"showFractionHint":false,"answer":"{answers_a[0]}"}]},{"type":"gapfill","prompt":"{prompts[1]}  [[0]] \n<br/>  <br/>","marks":0,"steps":[],"stepsPenalty":0,"showCorrectAnswer":true,"showFeedbackIcon":true,"useCustomName":true,"customName":"","variableReplacements":[],"variableReplacementStrategy":"originalfirst","adaptiveMarkingPenalty":0,"customMarkingAlgorithm":"","extendBaseMarkingAlgorithm":true,"sortAnswers":false,"gaps":[{"type":"numberentry","prompt":"","marks":1,"steps":[],"stepsPenalty":0,"showCorrectAnswer":true,"showFeedbackIcon":true,"useCustomName":true,"customName":"","variableReplacements":[],"variableReplacementStrategy":"originalfirst","adaptiveMarkingPenalty":0,"customMarkingAlgorithm":"","extendBaseMarkingAlgorithm":true,"correctAnswerFraction":false,"correctAnswerStyle":"plain","allowFractions":true,"notationStyles":["plain"],"mustBeReduced":false,"mustBeReducedPC":0,"precisionType":"none","precision":0,"precisionPartialCredit":0,"precisionMessage":"","strictPrecision":true,"showPrecisionHint":true,"showFractionHint":false,"answer":"{answers_a[1]}"}]},{"type":"gapfill","prompt":"{prompts[2]}  [[0]] \n<br/>  <br/>","marks":0,"steps":[],"stepsPenalty":0,"showCorrectAnswer":true,"showFeedbackIcon":true,"useCustomName":true,"customName":"","variableReplacements":[],"variableReplacementStrategy":"originalfirst","adaptiveMarkingPenalty":0,"customMarkingAlgorithm":"","extendBaseMarkingAlgorithm":true,"sortAnswers":false,"gaps":[{"type":"numberentry","prompt":"","marks":1,"steps":[],"stepsPenalty":0,"showCorrectAnswer":true,"showFeedbackIcon":true,"useCustomName":true,"customName":"","variableReplacements":[],"variableReplacementStrategy":"originalfirst","adaptiveMarkingPenalty":0,"customMarkingAlgorithm":"","extendBaseMarkingAlgorithm":true,"correctAnswerFraction":false,"correctAnswerStyle":"plain","allowFractions":true,"notationStyles":["plain"],"mustBeReduced":false,"mustBeReducedPC":0,"precisionType":"none","precision":0,"precisionPartialCredit":0,"precisionMessage":"","strictPrecision":true,"showPrecisionHint":true,"showFractionHint":false,"answer":"{answers_a[2]}"}]}],"builtin_constants":{"e":true,"i":true,"pi,π":true},"constants":[],"variables":{"advices":{"name":"advices","definition":"join(reorder(advices_formulas,selection),'<br/>')","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"advices_formulas":{"name":"advices_formulas","definition":"[\"\\\\[ \\\\var{price_orig} \\\\cdot (1- \\\\frac{\\\\var{percent}}{100})    = \\\\var{price_now} \\\\]\",\"\\\\[ \\\\frac{\\\\var{price_now1}}{ 1- \\\\frac{\\\\var{percent1}}{100}}    = \\\\var{price_orig1} \\\\]\",\"\\\\[ 100 \\\\cdot (1 - \\\\frac{\\\\var{price_now2}}{\\\\var{price_orig2}}) = \\\\var{percent2} \\\\]\"]","description":"","templateType":"list of strings","group":"Ungrouped variables","can_override":false},"answers":{"name":"answers","definition":"[ {price_now}, {price_orig1}, {percent2} ]","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"answers_a":{"name":"answers_a","definition":"reorder(answers,selection)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"percent":{"name":"percent","definition":"random(20,30,40,50,60)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"percent1":{"name":"percent1","definition":"{percent","description":"","templateType":"string","group":"Ungrouped variables","can_override":false},"percent2":{"name":"percent2","definition":"{percent}","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"price_now":{"name":"price_now","definition":"price_orig*(1-percent/100)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"price_now1":{"name":"price_now1","definition":"price_orig1*(1-percent1/100)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"price_now2":{"name":"price_now2","definition":"price_orig2*(1-percent2/100)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"price_orig":{"name":"price_orig","definition":"random(15,20,25,30,35)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"price_orig1":{"name":"price_orig1","definition":"{price_orig}","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"price_orig2":{"name":"price_orig2","definition":"{price_orig}","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"prompts":{"name":"prompts","definition":"reorder(prompts_nl,selection)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"prompts_en":{"name":"prompts_en","definition":"[\"Oorspronkelijk kostte je paar schoenen $\\\\var{price_orig}$ euro.\\nJe krijgt nu echter $\\\\var{percent}$% korting.\\n<br/>\\nHoeveel kost je paar schoenen nog? \\n\",\"Je schoenen kosten nu nog $\\\\var{price_now1}$ euro, want er is een korting van $\\\\var{percent1}$% op.\\n<br/>\\nHoeveel kostte je paar schoenen oorspronkelijk?\\n\",\"Oorspronkelijk kostte je paar schoenen $\\\\var{price_orig2}$ euro.\\nNu moet je maar $\\\\var{price_now2}$ euro betalen.\\n<br/>\\nHoeveel procent korting heb je gekregen op je paar schoenen?\\n\"]","description":"","templateType":"list of strings","group":"Ungrouped variables","can_override":false},"prompts_nl":{"name":"prompts_nl","definition":"[\"Oorspronkelijk kostte je paar schoenen $\\\\var{price_orig}$ euro.\\nJe krijgt nu echter $\\\\var{percent}$% korting.\\n<br/>\\nHoeveel kost je paar schoenen nog? \\n\",\"Je schoenen kosten nu nog $\\\\var{price_now1}$ euro, want er is een korting van $\\\\var{percent1}$% op.\\n<br/>\\nHoeveel kostte je paar schoenen oorspronkelijk?\\n\",\"Oorspronkelijk kostte je paar schoenen $\\\\var{price_orig2}$ euro.\\nNu moet je maar $\\\\var{price_now2}$ euro betalen.\\n<br/>\\nHoeveel procent korting heb je gekregen op je paar schoenen?\\n\"]","description":"","templateType":"list of strings","group":"Ungrouped variables","can_override":false},"selection":{"name":"selection","definition":"shuffle( 0..2 )","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"xanswers":{"name":"xanswers","definition":"[\"{price_now}\",\"{price_orig}\",\"{percent}\"]","description":"","templateType":"list of strings","group":"Ungrouped variables","can_override":false}},"variablesTest":{"condition":"","maxRuns":100},"functions":{},"ungrouped_variables":["advices","advices_formulas","answers","answers_a","percent","percent1","percent2","price_now","price_now1","price_now2","price_orig","price_orig1","price_orig2","prompts","prompts_en","prompts_nl","selection","xanswers"],"variable_groups":[],"rulesets":{"std":["basic","unitFactor","unitPower","unitDenominator","zeroFactor","zeroTerm","zeroPower","zeroBase","constantsFirst","sqrtProduct","sqrtDivision","sqrtSquare","otherNumbers","simplifyFractions","trig","cancelTerms","cancelFactors","collectLikeFractions"]},"preamble":{"js":"","css":"tt { font-family:monospace !important; background:lightgray !important}\n.tt { font-family:monospace !important; background:lightgray !important}\nul.info { text-indent: -28px; margin-left: 28px;}\nh4 { display: inline;}    /* prevent newline after a), b), ... in parts; */\ntable { margin-left: unset; }     /* tables flush-left ... */\ntable.open td { padding-bottom: 4rem !important; }  /* bigger padding in template-OPENx */\nli.instruction {\n  /* background-color: lightgray; */\n  /* list-style-type: none; */\n  /* display: inline-block;   /* reset list-item */\n  font-size: smaller;\n  margin-left: 2rem;\n  color: gray;\n  list-style-position: outside;\n  /* border: gray solid; */\n  /* padding: 0.5rem;    */\n}\n"},"navigation":{"allowregen":true,"showfrontpage":false,"preventleave":false},"extensions":[],"resources":[],"tags":[],"custom_part_types":[]},{"name":"percentages-shoes.yaml","statement":"You want to buy a pair of shoes, but wait for the sales to start.\n<li class=\"instruction\">De vraag met slechts één 'part'</li>","advice":"{advices}","parts":[{"type":"gapfill","prompt":"{prompts[0]}  [[0]]","marks":0,"steps":[],"stepsPenalty":0,"showCorrectAnswer":true,"showFeedbackIcon":true,"useCustomName":true,"customName":"","variableReplacements":[],"variableReplacementStrategy":"originalfirst","adaptiveMarkingPenalty":0,"customMarkingAlgorithm":"","extendBaseMarkingAlgorithm":true,"sortAnswers":false,"gaps":[{"type":"numberentry","prompt":"","marks":1,"steps":[],"stepsPenalty":0,"showCorrectAnswer":true,"showFeedbackIcon":true,"useCustomName":true,"customName":"","variableReplacements":[],"variableReplacementStrategy":"originalfirst","adaptiveMarkingPenalty":0,"customMarkingAlgorithm":"","extendBaseMarkingAlgorithm":true,"correctAnswerFraction":false,"correctAnswerStyle":"plain","allowFractions":true,"notationStyles":["plain"],"mustBeReduced":false,"mustBeReducedPC":0,"precisionType":"none","precision":0,"precisionPartialCredit":0,"precisionMessage":"","strictPrecision":true,"showPrecisionHint":true,"showFractionHint":false,"answer":"{answers_a[0]}"}]}],"builtin_constants":{"e":true,"i":true,"pi,π":true},"constants":[],"variables":{"advices":{"name":"advices","definition":"join(reorder(advices_formulas,selection),'<br/>')","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"advices_formulas":{"name":"advices_formulas","definition":"[\"\\\\[ \\\\var{price_orig} \\\\cdot (1- \\\\frac{\\\\var{percent}}{100})    = \\\\var{price_now} \\\\]\",\"\\\\[ \\\\frac{\\\\var{price_now1}}{ 1- \\\\frac{\\\\var{percent1}}{100}}    = \\\\var{price_orig1} \\\\]\",\"\\\\[ 100 \\\\cdot (1 - \\\\frac{\\\\var{price_now2}}{\\\\var{price_orig2}}) = \\\\var{percent2} \\\\]\"]","description":"","templateType":"list of strings","group":"Ungrouped variables","can_override":false},"answers":{"name":"answers","definition":"[ {price_now}, {price_orig1}, {percent2} ]","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"answers_a":{"name":"answers_a","definition":"reorder(answers,selection)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"percent":{"name":"percent","definition":"random(20,30,40,50,60)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"percent1":{"name":"percent1","definition":"{percent","description":"","templateType":"string","group":"Ungrouped variables","can_override":false},"percent2":{"name":"percent2","definition":"{percent}","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"price_now":{"name":"price_now","definition":"price_orig*(1-percent/100)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"price_now1":{"name":"price_now1","definition":"price_orig1*(1-percent1/100)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"price_now2":{"name":"price_now2","definition":"price_orig2*(1-percent2/100)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"price_orig":{"name":"price_orig","definition":"random(15,20,25,30,35)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"price_orig1":{"name":"price_orig1","definition":"{price_orig}","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"price_orig2":{"name":"price_orig2","definition":"{price_orig}","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"prompts":{"name":"prompts","definition":"reorder(prompts_nl,selection)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"prompts_en":{"name":"prompts_en","definition":"[\"Oorspronkelijk kostte je paar schoenen $\\\\var{price_orig}$ euro.\\nJe krijgt nu echter $\\\\var{percent}$% korting.\\n<br/>\\nHoeveel kost je paar schoenen nog? \\n\",\"Je schoenen kosten nu nog $\\\\var{price_now1}$ euro, want er is een korting van $\\\\var{percent1}$% op.\\n<br/>\\nHoeveel kostte je paar schoenen oorspronkelijk?\\n\",\"Oorspronkelijk kostte je paar schoenen $\\\\var{price_orig2}$ euro.\\nNu moet je maar $\\\\var{price_now2}$ euro betalen.\\n<br/>\\nHoeveel procent korting heb je gekregen op je paar schoenen?\\n\"]","description":"","templateType":"list of strings","group":"Ungrouped variables","can_override":false},"prompts_nl":{"name":"prompts_nl","definition":"[\"Oorspronkelijk kostte je paar schoenen $\\\\var{price_orig}$ euro.\\nJe krijgt nu echter $\\\\var{percent}$% korting.\\n<br/>\\nHoeveel kost je paar schoenen nog? \\n\",\"Je schoenen kosten nu nog $\\\\var{price_now1}$ euro, want er is een korting van $\\\\var{percent1}$% op.\\n<br/>\\nHoeveel kostte je paar schoenen oorspronkelijk?\\n\",\"Oorspronkelijk kostte je paar schoenen $\\\\var{price_orig2}$ euro.\\nNu moet je maar $\\\\var{price_now2}$ euro betalen.\\n<br/>\\nHoeveel procent korting heb je gekregen op je paar schoenen?\\n\"]","description":"","templateType":"list of strings","group":"Ungrouped variables","can_override":false},"selection":{"name":"selection","definition":"shuffle( 0..2 )","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"xanswers":{"name":"xanswers","definition":"[\"{price_now}\",\"{price_orig}\",\"{percent}\"]","description":"","templateType":"list of strings","group":"Ungrouped variables","can_override":false}},"variablesTest":{"condition":"","maxRuns":100},"functions":{},"ungrouped_variables":["advices","advices_formulas","answers","answers_a","percent","percent1","percent2","price_now","price_now1","price_now2","price_orig","price_orig1","price_orig2","prompts","prompts_en","prompts_nl","selection","xanswers"],"variable_groups":[],"rulesets":{"std":["basic","unitFactor","unitPower","unitDenominator","zeroFactor","zeroTerm","zeroPower","zeroBase","constantsFirst","sqrtProduct","sqrtDivision","sqrtSquare","otherNumbers","simplifyFractions","trig","cancelTerms","cancelFactors","collectLikeFractions"]},"preamble":{"js":"","css":"tt { font-family:monospace !important; background:lightgray !important}\n.tt { font-family:monospace !important; background:lightgray !important}\nul.info { text-indent: -28px; margin-left: 28px;}\nh4 { display: inline;}    /* prevent newline after a), b), ... in parts; */\ntable { margin-left: unset; }     /* tables flush-left ... */\ntable.open td { padding-bottom: 4rem !important; }  /* bigger padding in template-OPENx */\nli.instruction {\n  /* background-color: lightgray; */\n  /* list-style-type: none; */\n  /* display: inline-block;   /* reset list-item */\n  font-size: smaller;\n  margin-left: 2rem;\n  color: gray;\n  list-style-position: outside;\n  /* border: gray solid; */\n  /* padding: 0.5rem;    */\n}\n"},"navigation":{"allowregen":true,"showfrontpage":false,"preventleave":false},"extensions":[],"resources":[],"tags":[],"custom_part_types":[]},{"name":"percentages-shoes.yaml","statement":"You want to buy a pair of shoes, but wait for the sales to start.\n<li class=\"instruction\">Zonder de eerste vraag (bereken de nieuwe prijs, erg eenvouidg).</li>","advice":"{advices}","parts":[{"type":"gapfill","prompt":"{prompts[0]}  [[0]]","marks":0,"steps":[],"stepsPenalty":0,"showCorrectAnswer":true,"showFeedbackIcon":true,"useCustomName":true,"customName":"","variableReplacements":[],"variableReplacementStrategy":"originalfirst","adaptiveMarkingPenalty":0,"customMarkingAlgorithm":"","extendBaseMarkingAlgorithm":true,"sortAnswers":false,"gaps":[{"type":"numberentry","prompt":"","marks":1,"steps":[],"stepsPenalty":0,"showCorrectAnswer":true,"showFeedbackIcon":true,"useCustomName":true,"customName":"","variableReplacements":[],"variableReplacementStrategy":"originalfirst","adaptiveMarkingPenalty":0,"customMarkingAlgorithm":"","extendBaseMarkingAlgorithm":true,"correctAnswerFraction":false,"correctAnswerStyle":"plain","allowFractions":true,"notationStyles":["plain"],"mustBeReduced":false,"mustBeReducedPC":0,"precisionType":"none","precision":0,"precisionPartialCredit":0,"precisionMessage":"","strictPrecision":true,"showPrecisionHint":true,"showFractionHint":false,"answer":"{answers_a[0]}"}]}],"builtin_constants":{"e":true,"i":true,"pi,π":true},"constants":[],"variables":{"advices":{"name":"advices","definition":"join(reorder(advices_formulas,selection),'<br/>')","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"advices_formulas":{"name":"advices_formulas","definition":"[\"\\\\[ \\\\var{price_orig} \\\\cdot (1- \\\\frac{\\\\var{percent}}{100})    = \\\\var{price_now} \\\\]\",\"\\\\[ \\\\frac{\\\\var{price_now1}}{ 1- \\\\frac{\\\\var{percent1}}{100}}    = \\\\var{price_orig1} \\\\]\",\"\\\\[ 100 \\\\cdot (1 - \\\\frac{\\\\var{price_now2}}{\\\\var{price_orig2}}) = \\\\var{percent2} \\\\]\"]","description":"","templateType":"list of strings","group":"Ungrouped variables","can_override":false},"answers":{"name":"answers","definition":"[ {price_now}, {price_orig1}, {percent2} ]","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"answers_a":{"name":"answers_a","definition":"reorder(answers,selection)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"percent":{"name":"percent","definition":"random(20,30,40,50,60)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"percent1":{"name":"percent1","definition":"{percent","description":"","templateType":"string","group":"Ungrouped variables","can_override":false},"percent2":{"name":"percent2","definition":"{percent}","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"price_now":{"name":"price_now","definition":"price_orig*(1-percent/100)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"price_now1":{"name":"price_now1","definition":"price_orig1*(1-percent1/100)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"price_now2":{"name":"price_now2","definition":"price_orig2*(1-percent2/100)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"price_orig":{"name":"price_orig","definition":"random(15,20,25,30,35)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"price_orig1":{"name":"price_orig1","definition":"{price_orig}","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"price_orig2":{"name":"price_orig2","definition":"{price_orig}","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"prompts":{"name":"prompts","definition":"reorder(prompts_nl,selection)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"prompts_en":{"name":"prompts_en","definition":"[\"Oorspronkelijk kostte je paar schoenen $\\\\var{price_orig}$ euro.\\nJe krijgt nu echter $\\\\var{percent}$% korting.\\n<br/>\\nHoeveel kost je paar schoenen nog? \\n\",\"Je schoenen kosten nu nog $\\\\var{price_now1}$ euro, want er is een korting van $\\\\var{percent1}$% op.\\n<br/>\\nHoeveel kostte je paar schoenen oorspronkelijk?\\n\",\"Oorspronkelijk kostte je paar schoenen $\\\\var{price_orig2}$ euro.\\nNu moet je maar $\\\\var{price_now2}$ euro betalen.\\n<br/>\\nHoeveel procent korting heb je gekregen op je paar schoenen?\\n\"]","description":"","templateType":"list of strings","group":"Ungrouped variables","can_override":false},"prompts_nl":{"name":"prompts_nl","definition":"[\"Oorspronkelijk kostte je paar schoenen $\\\\var{price_orig}$ euro.\\nJe krijgt nu echter $\\\\var{percent}$% korting.\\n<br/>\\nHoeveel kost je paar schoenen nog? \\n\",\"Je schoenen kosten nu nog $\\\\var{price_now1}$ euro, want er is een korting van $\\\\var{percent1}$% op.\\n<br/>\\nHoeveel kostte je paar schoenen oorspronkelijk?\\n\",\"Oorspronkelijk kostte je paar schoenen $\\\\var{price_orig2}$ euro.\\nNu moet je maar $\\\\var{price_now2}$ euro betalen.\\n<br/>\\nHoeveel procent korting heb je gekregen op je paar schoenen?\\n\"]","description":"","templateType":"list of strings","group":"Ungrouped variables","can_override":false},"selection":{"name":"selection","definition":"shuffle([1,2])","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"xanswers":{"name":"xanswers","definition":"[\"{price_now}\",\"{price_orig}\",\"{percent}\"]","description":"","templateType":"list of strings","group":"Ungrouped variables","can_override":false}},"variablesTest":{"condition":"","maxRuns":100},"functions":{},"ungrouped_variables":["advices","advices_formulas","answers","answers_a","percent","percent1","percent2","price_now","price_now1","price_now2","price_orig","price_orig1","price_orig2","prompts","prompts_en","prompts_nl","selection","xanswers"],"variable_groups":[],"rulesets":{"std":["basic","unitFactor","unitPower","unitDenominator","zeroFactor","zeroTerm","zeroPower","zeroBase","constantsFirst","sqrtProduct","sqrtDivision","sqrtSquare","otherNumbers","simplifyFractions","trig","cancelTerms","cancelFactors","collectLikeFractions"]},"preamble":{"js":"","css":"tt { font-family:monospace !important; background:lightgray !important}\n.tt { font-family:monospace !important; background:lightgray !important}\nul.info { text-indent: -28px; margin-left: 28px;}\nh4 { display: inline;}    /* prevent newline after a), b), ... in parts; */\ntable { margin-left: unset; }     /* tables flush-left ... */\ntable.open td { padding-bottom: 4rem !important; }  /* bigger padding in template-OPENx */\nli.instruction {\n  /* background-color: lightgray; */\n  /* list-style-type: none; */\n  /* display: inline-block;   /* reset list-item */\n  font-size: smaller;\n  margin-left: 2rem;\n  color: gray;\n  list-style-position: outside;\n  /* border: gray solid; */\n  /* padding: 0.5rem;    */\n}\n"},"navigation":{"allowregen":true,"showfrontpage":false,"preventleave":false},"extensions":[],"resources":[],"tags":[],"custom_part_types":[]},{"name":"percentages-shoes.yaml","statement":"You want to buy a pair of shoes, but wait for the sales to start.\n<li class=\"instruction\">Met telkens ANDERE prijzen/percentages.</li>","advice":"{advices}","parts":[{"type":"gapfill","prompt":"{prompts[0]}  [[0]]","marks":0,"steps":[],"stepsPenalty":0,"showCorrectAnswer":true,"showFeedbackIcon":true,"useCustomName":true,"customName":"","variableReplacements":[],"variableReplacementStrategy":"originalfirst","adaptiveMarkingPenalty":0,"customMarkingAlgorithm":"","extendBaseMarkingAlgorithm":true,"sortAnswers":false,"gaps":[{"type":"numberentry","prompt":"","marks":1,"steps":[],"stepsPenalty":0,"showCorrectAnswer":true,"showFeedbackIcon":true,"useCustomName":true,"customName":"","variableReplacements":[],"variableReplacementStrategy":"originalfirst","adaptiveMarkingPenalty":0,"customMarkingAlgorithm":"","extendBaseMarkingAlgorithm":true,"correctAnswerFraction":false,"correctAnswerStyle":"plain","allowFractions":true,"notationStyles":["plain"],"mustBeReduced":false,"mustBeReducedPC":0,"precisionType":"none","precision":0,"precisionPartialCredit":0,"precisionMessage":"","strictPrecision":true,"showPrecisionHint":true,"showFractionHint":false,"answer":"{answers_a[0]}"}]},{"type":"gapfill","prompt":"{prompts[1]}  [[0]] \n<br/>  <br/>","marks":0,"steps":[],"stepsPenalty":0,"showCorrectAnswer":true,"showFeedbackIcon":true,"useCustomName":true,"customName":"","variableReplacements":[],"variableReplacementStrategy":"originalfirst","adaptiveMarkingPenalty":0,"customMarkingAlgorithm":"","extendBaseMarkingAlgorithm":true,"sortAnswers":false,"gaps":[{"type":"numberentry","prompt":"","marks":1,"steps":[],"stepsPenalty":0,"showCorrectAnswer":true,"showFeedbackIcon":true,"useCustomName":true,"customName":"","variableReplacements":[],"variableReplacementStrategy":"originalfirst","adaptiveMarkingPenalty":0,"customMarkingAlgorithm":"","extendBaseMarkingAlgorithm":true,"correctAnswerFraction":false,"correctAnswerStyle":"plain","allowFractions":true,"notationStyles":["plain"],"mustBeReduced":false,"mustBeReducedPC":0,"precisionType":"none","precision":0,"precisionPartialCredit":0,"precisionMessage":"","strictPrecision":true,"showPrecisionHint":true,"showFractionHint":false,"answer":"{answers_a[1]}"}]},{"type":"gapfill","prompt":"{prompts[2]}  [[0]] \n<br/>  <br/>","marks":0,"steps":[],"stepsPenalty":0,"showCorrectAnswer":true,"showFeedbackIcon":true,"useCustomName":true,"customName":"","variableReplacements":[],"variableReplacementStrategy":"originalfirst","adaptiveMarkingPenalty":0,"customMarkingAlgorithm":"","extendBaseMarkingAlgorithm":true,"sortAnswers":false,"gaps":[{"type":"numberentry","prompt":"","marks":1,"steps":[],"stepsPenalty":0,"showCorrectAnswer":true,"showFeedbackIcon":true,"useCustomName":true,"customName":"","variableReplacements":[],"variableReplacementStrategy":"originalfirst","adaptiveMarkingPenalty":0,"customMarkingAlgorithm":"","extendBaseMarkingAlgorithm":true,"correctAnswerFraction":false,"correctAnswerStyle":"plain","allowFractions":true,"notationStyles":["plain"],"mustBeReduced":false,"mustBeReducedPC":0,"precisionType":"none","precision":0,"precisionPartialCredit":0,"precisionMessage":"","strictPrecision":true,"showPrecisionHint":true,"showFractionHint":false,"answer":"{answers_a[2]}"}]}],"builtin_constants":{"e":true,"i":true,"pi,π":true},"constants":[],"variables":{"advices":{"name":"advices","definition":"join(reorder(advices_formulas,selection),'<br/>')","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"advices_formulas":{"name":"advices_formulas","definition":"[\"\\\\[ \\\\var{price_orig} \\\\cdot (1- \\\\frac{\\\\var{percent}}{100})    = \\\\var{price_now} \\\\]\",\"\\\\[ \\\\frac{\\\\var{price_now1}}{ 1- \\\\frac{\\\\var{percent1}}{100}}    = \\\\var{price_orig1} \\\\]\",\"\\\\[ 100 \\\\cdot (1 - \\\\frac{\\\\var{price_now2}}{\\\\var{price_orig2}}) = \\\\var{percent2} \\\\]\"]","description":"","templateType":"list of strings","group":"Ungrouped variables","can_override":false},"answers":{"name":"answers","definition":"[ {price_now}, {price_orig1}, {percent2} ]","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"answers_a":{"name":"answers_a","definition":"reorder(answers,selection)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"percent":{"name":"percent","definition":"random(20,30,40,50,60)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"percent1":{"name":"percent1","definition":"{percent","description":"","templateType":"string","group":"Ungrouped variables","can_override":false},"percent2":{"name":"percent2","definition":"random(5,10,20,30,50)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"price_now":{"name":"price_now","definition":"price_orig*(1-percent/100)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"price_now1":{"name":"price_now1","definition":"price_orig1*(1-percent1/100)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"price_now2":{"name":"price_now2","definition":"price_orig2*(1-percent2/100)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"price_orig":{"name":"price_orig","definition":"random(15,20,25,30,35)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"price_orig1":{"name":"price_orig1","definition":"random(15,20,25,30,35)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"price_orig2":{"name":"price_orig2","definition":"{price_orig}","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"prompts":{"name":"prompts","definition":"reorder(prompts_nl,selection)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"prompts_en":{"name":"prompts_en","definition":"[\"Oorspronkelijk kostte je paar schoenen $\\\\var{price_orig}$ euro.\\nJe krijgt nu echter $\\\\var{percent}$% korting.\\n<br/>\\nHoeveel kost je paar schoenen nog? \\n\",\"Je schoenen kosten nu nog $\\\\var{price_now1}$ euro, want er is een korting van $\\\\var{percent1}$% op.\\n<br/>\\nHoeveel kostte je paar schoenen oorspronkelijk?\\n\",\"Oorspronkelijk kostte je paar schoenen $\\\\var{price_orig2}$ euro.\\nNu moet je maar $\\\\var{price_now2}$ euro betalen.\\n<br/>\\nHoeveel procent korting heb je gekregen op je paar schoenen?\\n\"]","description":"","templateType":"list of strings","group":"Ungrouped variables","can_override":false},"prompts_nl":{"name":"prompts_nl","definition":"[\"Oorspronkelijk kostte je paar schoenen $\\\\var{price_orig}$ euro.\\nJe krijgt nu echter $\\\\var{percent}$% korting.\\n<br/>\\nHoeveel kost je paar schoenen nog? \\n\",\"Je schoenen kosten nu nog $\\\\var{price_now1}$ euro, want er is een korting van $\\\\var{percent1}$% op.\\n<br/>\\nHoeveel kostte je paar schoenen oorspronkelijk?\\n\",\"Oorspronkelijk kostte je paar schoenen $\\\\var{price_orig2}$ euro.\\nNu moet je maar $\\\\var{price_now2}$ euro betalen.\\n<br/>\\nHoeveel procent korting heb je gekregen op je paar schoenen?\\n\"]","description":"","templateType":"list of strings","group":"Ungrouped variables","can_override":false},"selection":{"name":"selection","definition":"shuffle( 0..2 )","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"xanswers":{"name":"xanswers","definition":"[\"{price_now}\",\"{price_orig}\",\"{percent}\"]","description":"","templateType":"list of strings","group":"Ungrouped variables","can_override":false}},"variablesTest":{"condition":"","maxRuns":100},"functions":{},"ungrouped_variables":["advices","advices_formulas","answers","answers_a","percent","percent1","percent2","price_now","price_now1","price_now2","price_orig","price_orig1","price_orig2","prompts","prompts_en","prompts_nl","selection","xanswers"],"variable_groups":[],"rulesets":{"std":["basic","unitFactor","unitPower","unitDenominator","zeroFactor","zeroTerm","zeroPower","zeroBase","constantsFirst","sqrtProduct","sqrtDivision","sqrtSquare","otherNumbers","simplifyFractions","trig","cancelTerms","cancelFactors","collectLikeFractions"]},"preamble":{"js":"","css":"tt { font-family:monospace !important; background:lightgray !important}\n.tt { font-family:monospace !important; background:lightgray !important}\nul.info { text-indent: -28px; margin-left: 28px;}\nh4 { display: inline;}    /* prevent newline after a), b), ... in parts; */\ntable { margin-left: unset; }     /* tables flush-left ... */\ntable.open td { padding-bottom: 4rem !important; }  /* bigger padding in template-OPENx */\nli.instruction {\n  /* background-color: lightgray; */\n  /* list-style-type: none; */\n  /* display: inline-block;   /* reset list-item */\n  font-size: smaller;\n  margin-left: 2rem;\n  color: gray;\n  list-style-position: outside;\n  /* border: gray solid; */\n  /* padding: 0.5rem;    */\n}\n"},"navigation":{"allowregen":true,"showfrontpage":false,"preventleave":false},"extensions":[],"resources":[],"tags":[],"custom_part_types":[]},{"name":"percentages-shoes.yaml","statement":"You want to buy a pair of shoes, but wait for the sales to start.\n<li class=\"instruction\">Met meer willekeurige prijzen.</li>","advice":"{advices}","parts":[{"type":"gapfill","prompt":"{prompts[0]}  [[0]]","marks":0,"steps":[],"stepsPenalty":0,"showCorrectAnswer":true,"showFeedbackIcon":true,"useCustomName":true,"customName":"","variableReplacements":[],"variableReplacementStrategy":"originalfirst","adaptiveMarkingPenalty":0,"customMarkingAlgorithm":"","extendBaseMarkingAlgorithm":true,"sortAnswers":false,"gaps":[{"type":"numberentry","prompt":"","marks":1,"steps":[],"stepsPenalty":0,"showCorrectAnswer":true,"showFeedbackIcon":true,"useCustomName":true,"customName":"","variableReplacements":[],"variableReplacementStrategy":"originalfirst","adaptiveMarkingPenalty":0,"customMarkingAlgorithm":"","extendBaseMarkingAlgorithm":true,"correctAnswerFraction":false,"correctAnswerStyle":"plain","allowFractions":true,"notationStyles":["plain"],"mustBeReduced":false,"mustBeReducedPC":0,"precisionType":"none","precision":0,"precisionPartialCredit":0,"precisionMessage":"","strictPrecision":true,"showPrecisionHint":true,"showFractionHint":false,"answer":"{answers_a[0]}"}]},{"type":"gapfill","prompt":"{prompts[1]}  [[0]] \n<br/>  <br/>","marks":0,"steps":[],"stepsPenalty":0,"showCorrectAnswer":true,"showFeedbackIcon":true,"useCustomName":true,"customName":"","variableReplacements":[],"variableReplacementStrategy":"originalfirst","adaptiveMarkingPenalty":0,"customMarkingAlgorithm":"","extendBaseMarkingAlgorithm":true,"sortAnswers":false,"gaps":[{"type":"numberentry","prompt":"","marks":1,"steps":[],"stepsPenalty":0,"showCorrectAnswer":true,"showFeedbackIcon":true,"useCustomName":true,"customName":"","variableReplacements":[],"variableReplacementStrategy":"originalfirst","adaptiveMarkingPenalty":0,"customMarkingAlgorithm":"","extendBaseMarkingAlgorithm":true,"correctAnswerFraction":false,"correctAnswerStyle":"plain","allowFractions":true,"notationStyles":["plain"],"mustBeReduced":false,"mustBeReducedPC":0,"precisionType":"none","precision":0,"precisionPartialCredit":0,"precisionMessage":"","strictPrecision":true,"showPrecisionHint":true,"showFractionHint":false,"answer":"{answers_a[1]}"}]},{"type":"gapfill","prompt":"{prompts[2]}  [[0]] \n<br/>  <br/>","marks":0,"steps":[],"stepsPenalty":0,"showCorrectAnswer":true,"showFeedbackIcon":true,"useCustomName":true,"customName":"","variableReplacements":[],"variableReplacementStrategy":"originalfirst","adaptiveMarkingPenalty":0,"customMarkingAlgorithm":"","extendBaseMarkingAlgorithm":true,"sortAnswers":false,"gaps":[{"type":"numberentry","prompt":"","marks":1,"steps":[],"stepsPenalty":0,"showCorrectAnswer":true,"showFeedbackIcon":true,"useCustomName":true,"customName":"","variableReplacements":[],"variableReplacementStrategy":"originalfirst","adaptiveMarkingPenalty":0,"customMarkingAlgorithm":"","extendBaseMarkingAlgorithm":true,"correctAnswerFraction":false,"correctAnswerStyle":"plain","allowFractions":true,"notationStyles":["plain"],"mustBeReduced":false,"mustBeReducedPC":0,"precisionType":"none","precision":0,"precisionPartialCredit":0,"precisionMessage":"","strictPrecision":true,"showPrecisionHint":true,"showFractionHint":false,"answer":"{answers_a[2]}"}]}],"builtin_constants":{"e":true,"i":true,"pi,π":true},"constants":[],"variables":{"advices":{"name":"advices","definition":"join(reorder(advices_formulas,selection),'<br/>')","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"advices_formulas":{"name":"advices_formulas","definition":"[\"\\\\[ \\\\var{price_orig} \\\\cdot (1- \\\\frac{\\\\var{percent}}{100})    = \\\\var{price_now} \\\\]\",\"\\\\[ \\\\frac{\\\\var{price_now1}}{ 1- \\\\frac{\\\\var{percent1}}{100}}    = \\\\var{price_orig1} \\\\]\",\"\\\\[ 100 \\\\cdot (1 - \\\\frac{\\\\var{price_now2}}{\\\\var{price_orig2}}) = \\\\var{percent2} \\\\]\"]","description":"","templateType":"list of strings","group":"Ungrouped variables","can_override":false},"answers":{"name":"answers","definition":"[ {price_now}, {price_orig1}, {percent2} ]","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"answers_a":{"name":"answers_a","definition":"reorder(answers,selection)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"percent":{"name":"percent","definition":"random(5,10,20,30,50)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"percent1":{"name":"percent1","definition":"{percent","description":"","templateType":"string","group":"Ungrouped variables","can_override":false},"percent2":{"name":"percent2","definition":"{percent}","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"price_now":{"name":"price_now","definition":"price_orig*(1-percent/100)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"price_now1":{"name":"price_now1","definition":"price_orig1*(1-percent1/100)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"price_now2":{"name":"price_now2","definition":"price_orig2*(1-percent2/100)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"price_orig":{"name":"price_orig","definition":"random(30 .. 99#1)","description":"","templateType":"randrange","group":"Ungrouped variables","can_override":false},"price_orig1":{"name":"price_orig1","definition":"random(30 .. 99#1)","description":"","templateType":"randrange","group":"Ungrouped variables","can_override":false},"price_orig2":{"name":"price_orig2","definition":"random(30 .. 99#1)","description":"","templateType":"randrange","group":"Ungrouped variables","can_override":false},"prompts":{"name":"prompts","definition":"reorder(prompts_nl,selection)","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"prompts_en":{"name":"prompts_en","definition":"[\"Oorspronkelijk kostte je paar schoenen $\\\\var{price_orig}$ euro.\\nJe krijgt nu echter $\\\\var{percent}$% korting.\\n<br/>\\nHoeveel kost je paar schoenen nog? \\n\",\"Je schoenen kosten nu nog $\\\\var{price_now1}$ euro, want er is een korting van $\\\\var{percent1}$% op.\\n<br/>\\nHoeveel kostte je paar schoenen oorspronkelijk?\\n\",\"Oorspronkelijk kostte je paar schoenen $\\\\var{price_orig2}$ euro.\\nNu moet je maar $\\\\var{price_now2}$ euro betalen.\\n<br/>\\nHoeveel procent korting heb je gekregen op je paar schoenen?\\n\"]","description":"","templateType":"list of strings","group":"Ungrouped variables","can_override":false},"prompts_nl":{"name":"prompts_nl","definition":"[\"Oorspronkelijk kostte je paar schoenen $\\\\var{price_orig}$ euro.\\nJe krijgt nu echter $\\\\var{percent}$% korting.\\n<br/>\\nHoeveel kost je paar schoenen nog? \\n\",\"Je schoenen kosten nu nog $\\\\var{price_now1}$ euro, want er is een korting van $\\\\var{percent1}$% op.\\n<br/>\\nHoeveel kostte je paar schoenen oorspronkelijk?\\n\",\"Oorspronkelijk kostte je paar schoenen $\\\\var{price_orig2}$ euro.\\nNu moet je maar $\\\\var{price_now2}$ euro betalen.\\n<br/>\\nHoeveel procent korting heb je gekregen op je paar schoenen?\\n\"]","description":"","templateType":"list of strings","group":"Ungrouped variables","can_override":false},"selection":{"name":"selection","definition":"shuffle( 0..2 )","description":"","templateType":"anything","group":"Ungrouped variables","can_override":false},"xanswers":{"name":"xanswers","definition":"[\"{price_now}\",\"{price_orig}\",\"{percent}\"]","description":"","templateType":"list of strings","group":"Ungrouped variables","can_override":false}},"variablesTest":{"condition":"","maxRuns":100},"functions":{},"ungrouped_variables":["advices","advices_formulas","answers","answers_a","percent","percent1","percent2","price_now","price_now1","price_now2","price_orig","price_orig1","price_orig2","prompts","prompts_en","prompts_nl","selection","xanswers"],"variable_groups":[],"rulesets":{"std":["basic","unitFactor","unitPower","unitDenominator","zeroFactor","zeroTerm","zeroPower","zeroBase","constantsFirst","sqrtProduct","sqrtDivision","sqrtSquare","otherNumbers","simplifyFractions","trig","cancelTerms","cancelFactors","collectLikeFractions"]},"preamble":{"js":"","css":"tt { font-family:monospace !important; background:lightgray !important}\n.tt { font-family:monospace !important; background:lightgray !important}\nul.info { text-indent: -28px; margin-left: 28px;}\nh4 { display: inline;}    /* prevent newline after a), b), ... in parts; */\ntable { margin-left: unset; }     /* tables flush-left ... */\ntable.open td { padding-bottom: 4rem !important; }  /* bigger padding in template-OPENx */\nli.instruction {\n  /* background-color: lightgray; */\n  /* list-style-type: none; */\n  /* display: inline-block;   /* reset list-item */\n  font-size: smaller;\n  margin-left: 2rem;\n  color: gray;\n  list-style-position: outside;\n  /* border: gray solid; */\n  /* padding: 0.5rem;    */\n}\n"},"navigation":{"allowregen":true,"showfrontpage":false,"preventleave":false},"extensions":[],"resources":[],"tags":[],"custom_part_types":[]}]}]}