// Numbas version: exam_results_page_options
{"name": "exam matrices eerste versie", "metadata": {"description": "", "licence": "None specified"}, "duration": 0, "percentPass": 0, "showQuestionGroupNames": false, "shuffleQuestionGroups": false, "showstudentname": true, "question_groups": [{"name": "Group", "pickingStrategy": "all-ordered", "pickQuestions": 1, "questionNames": ["", "", "", "", "", "", "", "", ""], "variable_overrides": [[], [], [], [], [], [], [], [], []], "questions": [{"name": "Find the Determinant of a 3 x 3 Matrix", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false}, "contributors": [{"name": "Newcastle University Mathematics and Statistics", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/697/"}], "variable_groups": [{"variables": ["a"], "name": "Unnamed group"}], "variables": {"a": {"templateType": "anything", "group": "Unnamed group", "definition": "matrix([ [a11,a12,a13],[a21,a22,a23],[a31,a32,a33] ])", "description": "", "name": "a"}, "a13": {"templateType": "anything", "group": "Ungrouped variables", "definition": "random(-5..5)", "description": "<p>Matrix element</p>", "name": "a13"}, "a22": {"templateType": "anything", "group": "Ungrouped variables", "definition": "random(-6..6 except 0)", "description": "<p>Matrix element</p>", "name": "a22"}, "a31": {"templateType": "anything", "group": "Ungrouped variables", "definition": "random(-5..5)", "description": "<p>Matrix element</p>", "name": "a31"}, "a23": {"templateType": "anything", "group": "Ungrouped variables", "definition": "random(-5..5)", "description": "<p>Matrix element</p>", "name": "a23"}, "m3": {"templateType": "anything", "group": "Ungrouped variables", "definition": "(a21*a32)-(a22*a31)", "description": "<p>Submatrix</p>", "name": "m3"}, "m1": {"templateType": "anything", "group": "Ungrouped variables", "definition": "(a22*a33)-(a23*a32)", "description": "<p>Submatrix</p>", "name": "m1"}, "m2": {"templateType": "anything", "group": "Ungrouped variables", "definition": "(a21*a33)-(a23*a31)", "description": "<p>Submatrix</p>", "name": "m2"}, "a32": {"templateType": "anything", "group": "Ungrouped variables", "definition": "random(-5..5)", "description": "<p>Matrix element</p>", "name": "a32"}, "a12": {"templateType": "anything", "group": "Ungrouped variables", "definition": "random(-5..5)", "description": "<p>Matrix element</p>", "name": "a12"}, "a11": {"templateType": "anything", "group": "Ungrouped variables", "definition": "random(-6..6 except 0)", "description": "<p>Matrix element</p>", "name": "a11"}, "a33": {"templateType": "anything", "group": "Ungrouped variables", "definition": "random(-6..6 except 0)", "description": "<p>Matrix element</p>", "name": "a33"}, "a21": {"templateType": "anything", "group": "Ungrouped variables", "definition": "random(-5..5)", "description": "<p>Matrix element</p>", "name": "a21"}}, "ungrouped_variables": ["a11", "a12", "a13", "a21", "a22", "a23", "a31", "a32", "a33", "m1", "m2", "m3"], "question_groups": [{"pickingStrategy": "all-ordered", "questions": [], "name": "", "pickQuestions": 0}], "functions": {}, "showQuestionGroupNames": false, "parts": [{"prompt": "<p>Calculate the determinant of the matrix.</p>\n<p>$\\operatorname{det}\\left( \\mathrm{A}\\right) = $&nbsp;[[0]]</p>\n<p></p>", "scripts": {}, "gaps": [{"correctAnswerFraction": false, "showCorrectAnswer": true, "scripts": {}, "allowFractions": false, "type": "numberentry", "maxValue": "det(a)", "minValue": "det(a)", "variableReplacementStrategy": "originalfirst", "variableReplacements": [], "marks": 0.5, "showPrecisionHint": false}], "type": "gapfill", "showCorrectAnswer": true, "variableReplacementStrategy": "originalfirst", "variableReplacements": [], "marks": 0}], "statement": "<p>Consider the&nbsp;$3 \\times 3$ matrix,</p>\n<p>\\begin{align} \\mathrm{A} &amp;= \\var{a} \\end{align}</p>", "tags": ["determinant of a matrix", "MAS1602", "mas1602", "matrices", "matrix"], "rulesets": {"std": ["all", "fractionNumbers", "!collectNumbers"]}, "preamble": {"css": "", "js": ""}, "type": "question", "metadata": {"notes": "", "licence": "Creative Commons Attribution 4.0 International", "description": "<p>Find the determinant of a&nbsp;$3 \\times 3$ matrix.</p>"}, "variablesTest": {"condition": "", "maxRuns": 100}, "advice": "<h4>Determinant of a $3 \\times 3$ matrix</h4>\n<p>The determinant of a matrix $\\mathrm{M} = \\begin{pmatrix} a&amp;b&amp;c \\\\ d&amp;e&amp;f \\\\ g&amp;h&amp;i \\end{pmatrix}$&nbsp;can be calculated&nbsp;by using cofactor expansion. Expanding along the first row,</p>\n<p>\\[ \\det\\left(\\mathrm{M}\\right) = a \\cdot \\det&nbsp;\\begin{pmatrix} e&amp;f \\\\ h&amp;i \\end{pmatrix}- b&nbsp;\\cdot \\det&nbsp;\\begin{pmatrix} d&amp;f \\\\ g&amp;i \\end{pmatrix} + c&nbsp;\\cdot \\det&nbsp;\\begin{pmatrix} d&amp;e \\\\ g&amp;h \\end{pmatrix}\\]</p>\n<p>Thus for our example we have:</p>\n<p>\\[\\begin{align} \\det&nbsp;\\begin{pmatrix} e&amp;f \\\\ h&amp;i \\end{pmatrix} &amp;= \\simplify[]{({a22}*{a33})-({a23}*{a32})&nbsp;= {m1}}&nbsp;\\\\ \\det&nbsp;\\begin{pmatrix} d&amp;f \\\\ g&amp;i \\end{pmatrix} &amp;= \\simplify[]{({a21}*{a33})-({a23}*{a31})&nbsp;=&nbsp;{m2}} \\\\ \\det&nbsp;\\begin{pmatrix} d&amp;e \\\\ g&amp;h \\end{pmatrix} &amp;=\\simplify[]{ ({a21}*{a32})-({a22}*{a31})&nbsp;={m3}} &nbsp;\\end{align}\\]</p>\n<p>and so</p>\n<p>\\[\\begin{align} \\det\\left(\\mathrm{A}\\right) = (\\simplify[]{{a11}*{m1}})-(\\simplify[]{{a12}*{m2}})+(\\simplify[]{{a13}*{m3}}) = \\simplify[]{{det(a)}} &nbsp;\\end{align}\\]</p>"}, {"name": "Multiply 2x2, 2x1 and 1x2 matrices", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false}, "contributors": [{"name": "Newcastle University Mathematics and Statistics", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/697/"}], "variable_groups": [], "variables": {"B": {"templateType": "anything", "group": "Ungrouped variables", "definition": "matrix([random(1..4),random(-1,1,2)],[random(-1,0,1),random(-4,-3,-2,-1,1,2,3,5)])", "description": "", "name": "B"}, "w": {"templateType": "anything", "group": "Ungrouped variables", "definition": "matrix([[random(-5..5),random(-7..7 except 0)]])", "description": "", "name": "w"}, "v": {"templateType": "anything", "group": "Ungrouped variables", "definition": "matrix([[random(-2..2)],[random(1..9)]])", "description": "", "name": "v"}}, "ungrouped_variables": ["B", "w", "v"], "question_groups": [{"pickingStrategy": "all-ordered", "questions": [], "name": "", "pickQuestions": 0}], "functions": {}, "showQuestionGroupNames": false, "parts": [{"scripts": {}, "gaps": [{"allowFractions": false, "correctAnswer": "w*v", "markPerCell": false, "allowResize": false, "correctAnswerFractions": false, "numRows": 1, "scripts": {}, "type": "matrix", "numColumns": 1, "tolerance": 0, "showCorrectAnswer": true, "marks": 1}], "type": "gapfill", "prompt": "<p>$\\mathbf{w} \\; \\mathbf{v} = $&nbsp;[[0]]</p>", "showCorrectAnswer": true, "marks": 0}, {"scripts": {}, "gaps": [{"allowFractions": false, "correctAnswer": "w*B", "markPerCell": false, "allowResize": false, "correctAnswerFractions": false, "numRows": 1, "scripts": {}, "type": "matrix", "numColumns": "2", "tolerance": 0, "showCorrectAnswer": true, "marks": 1}], "type": "gapfill", "prompt": "<p>$\\mathbf{w} \\mathbf{B} = $&nbsp;[[0]]</p>", "showCorrectAnswer": true, "marks": 0}, {"scripts": {}, "gaps": [{"allowFractions": false, "correctAnswer": "B*v", "markPerCell": false, "allowResize": false, "correctAnswerFractions": false, "numRows": "2", "scripts": {}, "type": "matrix", "numColumns": 1, "tolerance": 0, "showCorrectAnswer": true, "marks": 1}], "type": "gapfill", "prompt": "<p>$\\mathbf{B} \\mathbf{v} = $&nbsp;[[0]]</p>", "showCorrectAnswer": true, "marks": 0}], "statement": "<p>Answer the following questions on matrices.</p>\n<p>Let</p>\n<p>\\begin{align} \\mathbf{v} &amp;= \\var{v}, &amp; \\mathbf{B} &amp;= \\var{B}, &amp; \\mathbf{w} &amp;= \\var{w}&nbsp;\\end{align}</p>", "tags": ["checked2015", "MAS1602", "matrices", "matrix", "matrix multiplication", "matrix product", "multiplication of matrices", "multiply matrix", "product of matrices", "tested1"], "rulesets": {"std": ["all", "!collectNumbers", "fractionNumbers", "!noLeadingMinus"]}, "preamble": {"css": "", "js": ""}, "type": "question", "metadata": {"notes": "<p><strong>10/07/2012:</strong></p>\n<p>Added tags.<strong><br /></strong></p>\n<p>Perhaps it would be worthwhile restraining all components of the three matrices to be non zero?</p>\n<p>Question appears to be working correctly.</p>\n<p><strong>24/12/2012:</strong></p>\n<p>Calculation checked, OK. Added tested1 tag.</p>", "licence": "Creative Commons Attribution 4.0 International", "description": "<p>Very elementary matrix multiplication.&nbsp;</p>"}, "variablesTest": {"condition": "", "maxRuns": 100}, "advice": "<h4>a)</h4>\n<p>\\begin{align} <br/>\\mathbf{w} \\mathbf{v} &amp;= \\simplify{{w}*{v}}&nbsp;\\\\ <br/>&amp;= \\simplify[]{{w[0][0]}*{v[0][0]}+{w[0][1]}*{v[1][0]}} \\\\ <br/>&amp;= \\var{w*v} <br/>\\end{align}</p>\n<h4>b)</h4>\n<p>Since $\\mathbf{w}$ is a $1 \\times 2$ matrix and has the same number of columns, $2$, as the rows in the $2 \\times 2$ matrix $\\mathbf{B}$ we can multiply and form&nbsp;$\\mathbf{w} \\mathbf{B}$ which is a $1 \\times 2$ row vector:</p>\n<p>\\begin{align}<br/>\\mathbf{wB} &amp;= \\simplify{{w}*{B}} \\\\<br/>&amp;= \\begin{pmatrix} \\simplify[]{ {w[0][0]}*{B[0][0]} + {w[0][1]}*{B[1][0]} }, &amp; \\simplify[]{ {w[0][0]}*{B[0][1]} + {w[0][1]}*{B[1][1]} } \\end{pmatrix} \\\\<br/>&amp;= \\var{w*B}<br/>\\end{align}</p>\n<h4>c)</h4>\n<p>Since $\\mathbf{B}$ is a $2 \\times 2$ matrix and has the same number of columns, $2$, as the rows in $\\mathbf{v}$ we can form the matrix $\\mathbf{Bv}$ which will be a $2 \\times 1$ column vector.</p>\n<p>We find this vector to be:</p>\n<p>\\begin{align}<br/>\\mathbf{Bv} &amp;= \\simplify{{B}*{v}} \\\\<br/>&amp;= \\begin{pmatrix} \\simplify[]{ {B[0][0]}*{v[0][0]} + {B[0][1]}*{v[1][0]} } \\\\ \\simplify[]{ {B[1][0]}*{v[0][0]} + {B[1][1]}*{v[1][0]} } \\end{pmatrix} \\\\<br/>&amp;= \\var{B*v}<br/>\\end{align}</p>"}, {"name": "L1 - Matrix arithmetic combined", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false}, "contributors": [{"name": "Frank Doheny", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/789/"}, {"name": "John Steele", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2218/"}, {"name": "Josh Lim", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2990/"}, {"name": "Timur Zaripov", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/3272/"}], "statement": "<p>Given the matrix:</p>\n<p style=\"padding-left: 30px;\">\\(A=\\begin{pmatrix} \\var{a11}&amp; \\var{a12}\\\\ \\var{a21}&amp;\\var{a22}\\end{pmatrix}\\)</p>\n<p style=\"padding-left: 30px;\"></p>", "tags": [], "rulesets": {}, "variable_groups": [], "functions": {}, "advice": "<p style=\"padding-left: 30px;\">\\(A=\\begin{pmatrix} \\var{a11}&amp; \\var{a12}\\\\ \\var{a21}&amp;\\var{a22}\\end{pmatrix}\\)</p>\n<p style=\"padding-left: 30px;\">\\(A^2=\\begin{pmatrix} \\var{a11}&amp; \\var{a12}\\\\ \\var{a21}&amp;\\var{a22}\\end{pmatrix}\\begin{pmatrix} \\var{a11}&amp; \\var{a12}\\\\ \\var{a21}&amp;\\var{a22}\\end{pmatrix}\\)</p>\n<p style=\"padding-left: 30px;\">Remember multiplication of matrices is carried out by multiplying the rows of the first matrix by the columns of the second matrix.</p>\n<p style=\"padding-left: 30px;\">\\(A^2=\\begin{pmatrix} \\var{a11}&amp; \\var{a12}\\\\ \\var{a21}&amp;\\var{a22}\\end{pmatrix}\\begin{pmatrix} \\var{a11}&amp; \\var{a12}\\\\ \\var{a21}&amp;\\var{a22}\\end{pmatrix}=\\begin{pmatrix}\\var{a11}*\\var{a11}+\\var{a12}*\\var{a21}&amp;\\var{a11}*\\var{a12}+\\var{a12}*\\var{a22}\\\\ \\var{a21}*\\var{a11}+\\var{a22}*\\var{a21}&amp;\\var{a21}*\\var{a12}+\\var{a22}*\\var{a22}\\end{pmatrix}\\)</p>\n<p style=\"padding-left: 30px;\">\\(A^2=\\begin{pmatrix}\\simplify{{a11}*{a11}+{a12}*{a21}}&amp;\\simplify{{a11}*{a12}+{a12}*{a22}}\\\\ \\simplify{{a21}*{a11}+{a22}*{a21}}&amp;\\simplify{{a21}*{a12}+{a22}*{a22}}\\end{pmatrix}\\)</p>\n<p style=\"padding-left: 30px;\">\\(\\var{k1}A=\\begin{pmatrix} \\var{k1}*\\var{a11}&amp; \\var{k1}*\\var{a12}\\\\ \\var{k1}*\\var{a21}&amp;\\var{k1}*\\var{a22}\\end{pmatrix}=\\begin{pmatrix} \\simplify{{k1}*{a11}}&amp; \\simplify{{k1}*{a12}}\\\\ \\simplify{{k1}*{a21}}&amp;\\simplify{{k1}*{a22}}\\end{pmatrix}\\)</p>\n<p style=\"padding-left: 30px;\">\\(\\left(A^2+\\var{k1}A+\\var{k2}I\\right)^t=\\left(\\begin{pmatrix}\\simplify{{a11}*{a11}+{a12}*{a21}}&amp;\\simplify{{a11}*{a12}+{a12}*{a22}}\\\\ \\simplify{{a21}*{a11}+{a22}*{a21}}&amp;\\simplify{{a21}*{a12}+{a22}*{a22}}\\end{pmatrix}+\\begin{pmatrix} \\simplify{{k1}*{a11}}&amp; \\simplify{{k1}*{a12}}\\\\ \\simplify{{k1}*{a21}}&amp;\\simplify{{k1}*{a22}}\\end{pmatrix}+\\begin{pmatrix} \\var{k2}&amp;0\\\\0&amp;\\var{k2}\\end{pmatrix}\\right)^t\\)</p>\n<p style=\"padding-left: 30px;\">\\(\\left(A^2+\\var{k1}A+\\var{k2}I\\right)^t=\\begin{pmatrix}\\simplify{{a11}*{a11}+{a12}*{a21}+{k1}{a11}+{k2}}&amp;\\simplify{{a11}*{a12}+{a12}*{a22}+{k1}*{a12}}\\\\ \\simplify{{a21}*{a11}+{a22}*{a21}+{k1}*{a21}}&amp;\\simplify{{a21}*{a12}+{a22}*{a22}+{k1}*{a22}+{k2}}\\end{pmatrix}^t\\)</p>\n<p style=\"padding-left: 30px;\">\\(\\left(A^2+\\var{k1}A+\\var{k2}I\\right)^t=\\begin{pmatrix}\\simplify{{a11}*{a11}+{a12}*{a21}+{k1}{a11}+{k2}}&amp;\\simplify{{a21}*{a11}+{a22}*{a21}+{k1}*{a21}}\\\\ \\simplify{{a11}*{a12}+{a12}*{a22}+{k1}*{a12}}&amp;\\simplify{{a21}*{a12}+{a22}*{a22}+{k1}*{a22}+{k2}}\\end{pmatrix}\\)</p>", "ungrouped_variables": ["a11", "a12", "a21", "a22", "k1", "k2", "b11", "b12", "b21", "b22"], "variablesTest": {"maxRuns": 100, "condition": ""}, "preamble": {"js": "", "css": ""}, "variables": {"a11": {"name": "a11", "group": "Ungrouped variables", "templateType": "number", "description": "", "definition": "-2"}, "b11": {"name": "b11", "group": "Ungrouped variables", "templateType": "anything", "description": "", "definition": "{a11}^2+{a12}*{a21}+{k1}*{a11}+{k2}"}, "b21": {"name": "b21", "group": "Ungrouped variables", "templateType": "anything", "description": "", "definition": "{a12}*{a11}+{a22}*{a12}+{k1}*{a12}"}, "b22": {"name": "b22", "group": "Ungrouped variables", "templateType": "anything", "description": "", "definition": "{a21}*{a12}+{a22}^2+{k1}*{a22}+{k2}"}, "b12": {"name": "b12", "group": "Ungrouped variables", "templateType": "anything", "description": "", "definition": "{a21}*{a11}+{a22}*{a21}+{k1}*{a21}"}, "a21": {"name": "a21", "group": "Ungrouped variables", "templateType": "number", "description": "", "definition": "0"}, "k2": {"name": "k2", "group": "Ungrouped variables", "templateType": "number", "description": "", "definition": "2"}, "k1": {"name": "k1", "group": "Ungrouped variables", "templateType": "number", "description": "", "definition": "2"}, "a12": {"name": "a12", "group": "Ungrouped variables", "templateType": "number", "description": "", "definition": "-1"}, "a22": {"name": "a22", "group": "Ungrouped variables", "templateType": "number", "description": "", "definition": "-2"}}, "metadata": {"licence": "Creative Commons Attribution-NonCommercial 4.0 International", "description": "<p>This question tests students knowledge of basic matrix arithmetic.</p>"}, "parts": [{"useCustomName": false, "adaptiveMarkingPenalty": 0, "prompt": "<p>Evaluate the following expression:</p>\n<p style=\"padding-left: 60px;\">\\(\\left(A^2+\\var{k1}A+\\var{k2}I\\right)^T\\) =&nbsp;[[0]]</p>", "showCorrectAnswer": true, "marks": 0, "customName": "", "gaps": [{"useCustomName": false, "markPerCell": false, "adaptiveMarkingPenalty": 0, "showCorrectAnswer": true, "tolerance": 0, "correctAnswer": "matrix([\n   [b11,b12],\n   [b21,b22]\n])    ", "numRows": "2", "marks": "5", "customName": "", "allowFractions": false, "customMarkingAlgorithm": "", "variableReplacementStrategy": "originalfirst", "variableReplacements": [], "correctAnswerFractions": false, "allowResize": false, "showFeedbackIcon": true, "scripts": {}, "numColumns": "2", "type": "matrix", "extendBaseMarkingAlgorithm": true, "unitTests": []}], "sortAnswers": false, "customMarkingAlgorithm": "", "variableReplacementStrategy": "originalfirst", "showFeedbackIcon": true, "variableReplacements": [], "scripts": {}, "type": "gapfill", "extendBaseMarkingAlgorithm": true, "unitTests": []}]}, {"name": "L1 - Matrices intro and definitions", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false}, "contributors": [{"name": "Marie Nicholson", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/1799/"}, {"name": "Timur Zaripov", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/3272/"}, {"name": "Maria Aneiros", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/3388/"}], "ungrouped_variables": [], "advice": "<p>a) &nbsp; &nbsp;&nbsp; Remember that the dimensions are $\\rm \\color{red}{rows} \\times \\color{blue}{columns}$, so you need to count the number or rwos and columns in the matrix and wrtite them in that order.</p>\n<p>b) &nbsp; &nbsp;&nbsp; Remember that the elements are in the form $A_{\\rm \\color{red}{ rows},\\color{blue}{columns}}$ where $A$ is the matrix.</p>\n<p>For example if we are looking for $a_{12}$ we look at the matrix $A= \\begin{pmatrix} \\var{a11} &amp;\\bf( \\underline{\\var{a12}}) \\\\ \\var{a21} &amp; \\var{a22}\\end{pmatrix}$ we want the the element on $\\rm \\color{red}{row ~ 1}$ and $\\rm \\color{blue}{column ~ 2}$ which in this case is $\\bf \\underline{\\var{a12}}$</p>\n<p></p>", "functions": {}, "variablesTest": {"condition": "", "maxRuns": 100}, "parts": [{"adaptiveMarkingPenalty": 0, "useCustomName": false, "scripts": {}, "gaps": [{"adaptiveMarkingPenalty": 0, "mustBeReducedPC": 0, "scripts": {}, "variableReplacementStrategy": "originalfirst", "marks": 1, "minValue": "{n1}", "showFractionHint": true, "mustBeReduced": false, "customMarkingAlgorithm": "", "customName": "", "maxValue": "{n1}", "useCustomName": false, "type": "numberentry", "unitTests": [], "correctAnswerStyle": "plain", "notationStyles": ["plain", "en", "si-en"], "extendBaseMarkingAlgorithm": true, "showFeedbackIcon": true, "showCorrectAnswer": true, "allowFractions": false, "correctAnswerFraction": false, "variableReplacements": []}, {"adaptiveMarkingPenalty": 0, "mustBeReducedPC": 0, "scripts": {}, "variableReplacementStrategy": "originalfirst", "marks": 1, "minValue": "{m1}", "showFractionHint": true, "mustBeReduced": false, "customMarkingAlgorithm": "", "customName": "", "maxValue": "{m1}", "useCustomName": false, "type": "numberentry", "unitTests": [], "correctAnswerStyle": "plain", "notationStyles": ["plain", "en", "si-en"], "extendBaseMarkingAlgorithm": true, "showFeedbackIcon": true, "showCorrectAnswer": true, "allowFractions": false, "correctAnswerFraction": false, "variableReplacements": []}, {"adaptiveMarkingPenalty": 0, "mustBeReducedPC": 0, "scripts": {}, "variableReplacementStrategy": "originalfirst", "marks": 1, "minValue": "{n2}", "showFractionHint": true, "mustBeReduced": false, "customMarkingAlgorithm": "", "customName": "", "maxValue": "{n2}", "useCustomName": false, "type": "numberentry", "unitTests": [], "correctAnswerStyle": "plain", "notationStyles": ["plain", "en", "si-en"], "extendBaseMarkingAlgorithm": true, "showFeedbackIcon": true, "showCorrectAnswer": true, "allowFractions": false, "correctAnswerFraction": false, "variableReplacements": []}, {"adaptiveMarkingPenalty": 0, "mustBeReducedPC": 0, "scripts": {}, "variableReplacementStrategy": "originalfirst", "marks": 1, "minValue": "{m2}", "showFractionHint": true, "mustBeReduced": false, "customMarkingAlgorithm": "", "customName": "", "maxValue": "{m2}", "useCustomName": false, "type": "numberentry", "unitTests": [], "correctAnswerStyle": "plain", "notationStyles": ["plain", "en", "si-en"], "extendBaseMarkingAlgorithm": true, "showFeedbackIcon": true, "showCorrectAnswer": true, "allowFractions": false, "correctAnswerFraction": false, "variableReplacements": []}, {"adaptiveMarkingPenalty": 0, "mustBeReducedPC": 0, "scripts": {}, "variableReplacementStrategy": "originalfirst", "marks": 1, "minValue": "{n3}", "showFractionHint": true, "mustBeReduced": false, "customMarkingAlgorithm": "", "customName": "", "maxValue": "{n3}", "useCustomName": false, "type": "numberentry", "unitTests": [], "correctAnswerStyle": "plain", "notationStyles": ["plain", "en", "si-en"], "extendBaseMarkingAlgorithm": true, "showFeedbackIcon": true, "showCorrectAnswer": true, "allowFractions": false, "correctAnswerFraction": false, "variableReplacements": []}, {"adaptiveMarkingPenalty": 0, "mustBeReducedPC": 0, "scripts": {}, "variableReplacementStrategy": "originalfirst", "marks": 1, "minValue": "{m3}", "showFractionHint": true, "mustBeReduced": false, "customMarkingAlgorithm": "", "customName": "", "maxValue": "{m3}", "useCustomName": false, "type": "numberentry", "unitTests": [], "correctAnswerStyle": "plain", "notationStyles": ["plain", "en", "si-en"], "extendBaseMarkingAlgorithm": true, "showFeedbackIcon": true, "showCorrectAnswer": true, "allowFractions": false, "correctAnswerFraction": false, "variableReplacements": []}], "variableReplacementStrategy": "originalfirst", "marks": 0, "extendBaseMarkingAlgorithm": true, "type": "gapfill", "unitTests": [], "showCorrectAnswer": true, "customMarkingAlgorithm": "", "customName": "", "sortAnswers": false, "prompt": "<p><strong>Example 1 -- </strong>Give the dimensions of the following matrices:</p>\n<p>$A=\\var{A}$ has dimensions [[0]]$ \\times $[[1]]</p>\n<p>$B=\\var{B}$ is a [[2]]$ \\times$[[3]] matrix</p>\n<p>$C=\\var{C}$ is of dimension [[4]]$ \\times$[[5]]</p>", "showFeedbackIcon": true, "variableReplacements": []}, {"adaptiveMarkingPenalty": 0, "useCustomName": false, "scripts": {}, "gaps": [{"adaptiveMarkingPenalty": 0, "mustBeReducedPC": 0, "scripts": {}, "variableReplacementStrategy": "originalfirst", "marks": 1, "minValue": "A[n1-1][m1-1]", "showFractionHint": true, "mustBeReduced": false, "customMarkingAlgorithm": "", "customName": "", "maxValue": "A[n1-1][m1-1]", "useCustomName": false, "type": "numberentry", "unitTests": [], "correctAnswerStyle": "plain", "notationStyles": ["plain", "en", "si-en"], "extendBaseMarkingAlgorithm": true, "showFeedbackIcon": true, "showCorrectAnswer": true, "allowFractions": false, "correctAnswerFraction": false, "variableReplacements": []}, {"adaptiveMarkingPenalty": 0, "mustBeReducedPC": 0, "scripts": {}, "variableReplacementStrategy": "originalfirst", "marks": 1, "minValue": "b[n2-1][m2-1]", "showFractionHint": true, "mustBeReduced": false, "customMarkingAlgorithm": "", "customName": "", "maxValue": "b[n2-1][m2-1]", "useCustomName": false, "type": "numberentry", "unitTests": [], "correctAnswerStyle": "plain", "notationStyles": ["plain", "en", "si-en"], "extendBaseMarkingAlgorithm": true, "showFeedbackIcon": true, "showCorrectAnswer": true, "allowFractions": false, "correctAnswerFraction": false, "variableReplacements": []}, {"adaptiveMarkingPenalty": 0, "mustBeReducedPC": 0, "scripts": {}, "variableReplacementStrategy": "originalfirst", "marks": 1, "minValue": "c[n3-1][m3-1]", "showFractionHint": true, "mustBeReduced": false, "customMarkingAlgorithm": "", "customName": "", "maxValue": "c[n3-1][m3-1]", "useCustomName": false, "type": "numberentry", "unitTests": [], "correctAnswerStyle": "plain", "notationStyles": ["plain", "en", "si-en"], "extendBaseMarkingAlgorithm": true, "showFeedbackIcon": true, "showCorrectAnswer": true, "allowFractions": false, "correctAnswerFraction": false, "variableReplacements": []}, {"adaptiveMarkingPenalty": 0, "mustBeReducedPC": 0, "scripts": {}, "variableReplacementStrategy": "originalfirst", "marks": 1, "minValue": "c[n4-1][m4-1]", "showFractionHint": true, "mustBeReduced": false, "customMarkingAlgorithm": "", "customName": "", "maxValue": "c[n4-1][m4-1]", "useCustomName": false, "type": "numberentry", "unitTests": [], "correctAnswerStyle": "plain", "notationStyles": ["plain", "en", "si-en"], "extendBaseMarkingAlgorithm": true, "showFeedbackIcon": true, "showCorrectAnswer": true, "allowFractions": false, "correctAnswerFraction": false, "variableReplacements": []}], "variableReplacementStrategy": "originalfirst", "marks": 0, "extendBaseMarkingAlgorithm": true, "type": "gapfill", "unitTests": [], "showCorrectAnswer": true, "customMarkingAlgorithm": "", "customName": "", "sortAnswers": false, "prompt": "<p><strong>Example 2 -- </strong>Give the values of the following elements of the matrices from example 1 above:</p>\n<p>$a_{\\var{n1}\\var{m1}}=$[[0]]</p>\n<p>$b_{\\var{n2}\\var{m2}}=$[[1]]</p>\n<p>$c_{\\var{n3}\\var{m3}}=$[[2]]</p>\n<p>$c_{\\var{n4}\\var{m4}}=$[[3]]</p>", "showFeedbackIcon": true, "variableReplacements": []}, {"adaptiveMarkingPenalty": 0, "showCellAnswerState": true, "scripts": {}, "variableReplacementStrategy": "originalfirst", "marks": 0, "distractors": ["", "", ""], "maxMarks": 0, "unitTests": [], "customMarkingAlgorithm": "", "choices": ["<p>$\\mathrm{rows} \\times \\mathrm{columns}$</p>", "<p>$\\mathrm{columns} \\times \\mathrm{rows}$</p>", "<p>Neither of the above.</p>"], "customName": "", "shuffleChoices": false, "useCustomName": false, "showFeedbackIcon": true, "displayColumns": "1", "type": "1_n_2", "displayType": "radiogroup", "extendBaseMarkingAlgorithm": true, "matrix": ["1", 0, 0], "showCorrectAnswer": true, "minMarks": 0, "prompt": "<p>If a matrix is described as being $p \\times q$ then the order is</p>", "variableReplacements": []}, {"adaptiveMarkingPenalty": 0, "showCellAnswerState": true, "scripts": {}, "variableReplacementStrategy": "originalfirst", "marks": 0, "warningType": "none", "unitTests": [], "showFeedbackIcon": true, "choices": ["$i$", "$j$"], "minAnswers": 0, "customName": "", "shuffleChoices": false, "customMarkingAlgorithm": "", "useCustomName": false, "type": "m_n_x", "displayType": "radiogroup", "extendBaseMarkingAlgorithm": true, "answers": ["Columns", "Rows", "Rows or columns interchangeably"], "maxAnswers": 0, "layout": {"type": "all", "expression": ""}, "matrix": [[0, "0.5", 0], ["0.5", 0, 0]], "showCorrectAnswer": true, "shuffleAnswers": false, "minMarks": 0, "maxMarks": 0, "prompt": "<p>If a matrix $A$ is described as having elements $a_{ij}$ then the variables on the left can be described as follows:</p>", "variableReplacements": []}], "rulesets": {}, "variables": {"n3": {"description": "", "templateType": "anything", "name": "n3", "definition": "random(2..4)", "group": "C"}, "n2": {"description": "", "templateType": "anything", "name": "n2", "definition": "random(2..4)", "group": "B"}, "m4": {"description": "", "templateType": "anything", "name": "m4", "definition": "random(1..(m3-1) except m3)", "group": "C2"}, "C": {"description": "", "templateType": "anything", "name": "C", "definition": "transpose(matrix(repeat(repeat(random(-9..9),n3),m3)))", "group": "C"}, "m1": {"description": "", "templateType": "anything", "name": "m1", "definition": "random(1..2)", "group": "A"}, "m2": {"description": "", "templateType": "anything", "name": "m2", "definition": "random(2..4)", "group": "B"}, "B": {"description": "", "templateType": "anything", "name": "B", "definition": "transpose(matrix(repeat(repeat(random(-9..9),n2),m2)))", "group": "B"}, "n4": {"description": "", "templateType": "anything", "name": "n4", "definition": "random(1..(n3-1) except n3)", "group": "C2"}, "n1": {"description": "", "templateType": "anything", "name": "n1", "definition": "random(1..2)", "group": "A"}, "A": {"description": "", "templateType": "anything", "name": "A", "definition": "transpose(matrix(repeat(repeat(random(-9..9),n1),m1)))", "group": "A"}, "m3": {"description": "", "templateType": "anything", "name": "m3", "definition": "random(2..4)", "group": "C"}}, "tags": [], "statement": "<p>Matrices are rectangular arrays of numbers arranged in rows and columns. &nbsp;For example</p>\n<p>$\\begin{pmatrix} 1&amp;2\\\\ 4&amp;5\\\\ \\end{pmatrix} , &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp;\\begin{pmatrix} 0&amp;-1\\ &amp;3\\\\2&amp;4&amp;-2 \\end{pmatrix}$</p>\n<p>The rows run across the matrices while the columns run down the matrices. &nbsp;Thus in the first matrix above the numbers $\\begin{pmatrix}1&amp;2\\end{pmatrix}$ are in the first row while the numbers $\\begin{pmatrix}4&amp;5\\end{pmatrix}$ are in the second row and similarly the numbers $\\begin{pmatrix} 1\\\\ 4\\\\ \\end{pmatrix}$&nbsp;are in the first column while the numbers $\\begin{pmatrix} 2\\\\ 5\\\\ \\end{pmatrix}$&nbsp;are in the second column.</p>\n<p>A matrix with $\\bf&nbsp;m$ rows and $\\bf n$<strong>&nbsp;</strong>columns is called a matrix of <em>order</em> $\\bf&nbsp;m\\times n$ or <em>dimension</em> $\\bf m\\times&nbsp;n$&nbsp;(or an $\\bf m\\times&nbsp;n$ matrix for&nbsp;brevity).</p>\n<p>When working with matrices the positions of the numbers in the arrays are as important as the actual values of the numbers. &nbsp;Given a matrix called &nbsp;$A$<strong>&nbsp;</strong>the number in row $i$ and column $j$ is usually denoted $a_{ij}$ it is also sometimes called the $ij^{th}$ element of the matrix $A$.</p>", "metadata": {"description": "<p>a<sub>ij<sup>&nbsp; notation and definition of the order of a matrix.</sup></sub></p>", "licence": "Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International"}, "preamble": {"css": "", "js": ""}, "variable_groups": [{"variables": ["n1", "m1", "A"], "name": "A"}, {"variables": ["n2", "m2", "B"], "name": "B"}, {"variables": ["n3", "m3", "C"], "name": "C"}, {"variables": ["n4", "m4"], "name": "C2"}]}, {"name": "L1 - Matrix multiplication (selfstudy)", "extensions": ["stats"], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false}, "contributors": [{"name": "Bill Foster", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/6/"}, {"name": "Christian Lawson-Perfect", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/7/"}, {"name": "Timur Zaripov", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/3272/"}], "functions": {}, "rulesets": {"std": ["all", "!collectNumbers", "fractionNumbers", "!noLeadingMinus"]}, "metadata": {"licence": "Creative Commons Attribution 4.0 International", "description": "<p>Exercises in multiplying matrices.&nbsp;</p>"}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["ba21", "a21", "a22", "a23", "b23", "b22", "b21", "ba22", "w3", "w2", "w1", "ba11", "a11", "p1", "q1", "q3", "q2", "ab23", "ab22", "ab21", "s3", "s2", "s1", "a12", "s7", "s6", "s5", "s4", "b12", "b13", "b11", "c13", "c12", "c11", "a33", "a32", "a31", "ba32", "ba33", "a13", "ba31", "s10", "s9", "b32", "b31", "v1", "v2", "v3", "mae", "mad", "mag", "mah", "ba12", "ba13", "b33", "p2", "p3", "ba23", "ab11", "s8", "ab31", "ab32", "ab33", "ab12", "m", "mac", "n", "q", "p", "s", "r", "u", "w", "v", "y", "x", "z", "ab13"], "variable_groups": [], "advice": "", "variables": {"mah": {"group": "Ungrouped variables", "name": "mah", "templateType": "anything", "description": "", "definition": "matrix(repeat(repeat(random(-2..9),x),y))"}, "ab31": {"group": "Ungrouped variables", "name": "ab31", "templateType": "anything", "description": "", "definition": "a31*b11+a32*b21+a33*b31"}, "w3": {"group": "Ungrouped variables", "name": "w3", "templateType": "anything", "description": "", "definition": "random(-4..4)"}, "b33": {"group": "Ungrouped variables", "name": "b33", "templateType": "anything", "description": "", "definition": "random(-2..2)"}, "ba33": {"group": "Ungrouped variables", "name": "ba33", "templateType": "anything", "description": "", "definition": "b31*a13+b32*a23+b33*a33"}, "b31": {"group": "Ungrouped variables", "name": "b31", "templateType": "anything", "description": "", "definition": "random(1..3)"}, "q": {"group": "Ungrouped variables", "name": "q", "templateType": "anything", "description": "", "definition": "m+random(0,z)"}, "ab11": {"group": "Ungrouped variables", "name": "ab11", "templateType": "anything", "description": "", "definition": "a11*b11+a12*b21+a13*b31"}, "s7": {"group": "Ungrouped variables", "name": "s7", "templateType": "anything", "description": "", "definition": "if(u=y,0.5,-0.5)"}, "ab21": {"group": "Ungrouped variables", "name": "ab21", "templateType": "anything", "description": "", "definition": "a21*b11+a22*b21+a23*b31"}, "mac": {"group": "Ungrouped variables", "name": "mac", "templateType": "anything", "description": "", "definition": "matrix(repeat(repeat(random(-2..9),n),m))"}, "s3": {"group": "Ungrouped variables", "name": "s3", "templateType": "anything", "description": "", "definition": "if(s=5,0.5,-0.5)"}, "b32": {"group": "Ungrouped variables", "name": "b32", "templateType": "anything", "description": "", "definition": "random(-3..3)"}, "r": {"group": "Ungrouped variables", "name": "r", "templateType": "anything", "description": "", "definition": "random(1..3)"}, "v3": {"group": "Ungrouped variables", "name": "v3", "templateType": "anything", "description": "", "definition": "random(-5..5)"}, "u": {"group": "Ungrouped variables", "name": "u", "templateType": "anything", "description": "", "definition": "random(3..4)"}, "c11": {"group": "Ungrouped variables", "name": "c11", "templateType": "anything", "description": "", "definition": "random(-2..2)"}, "b23": {"group": "Ungrouped variables", "name": "b23", "templateType": "anything", "description": "", "definition": "random(1,2)"}, "a32": {"group": "Ungrouped variables", "name": "a32", "templateType": "anything", "description": "", "definition": "random(-4..4)"}, "v1": {"group": "Ungrouped variables", "name": "v1", "templateType": "anything", "description": "", "definition": "random(-3..3)"}, "b21": {"group": "Ungrouped variables", "name": "b21", "templateType": "anything", "description": "", "definition": "random(0..2)"}, "ba22": {"group": "Ungrouped variables", "name": "ba22", "templateType": "anything", "description": "", "definition": "b21*a12+b22*a22+b23*a32"}, "ab33": {"group": "Ungrouped variables", "name": "ab33", "templateType": "anything", "description": "", "definition": "a31*b13+a32*b23+a33*b33"}, "mag": {"group": "Ungrouped variables", "name": "mag", "templateType": "anything", "description": "", "definition": "matrix(repeat(repeat(random(-2..9),u),w))"}, "a12": {"group": "Ungrouped variables", "name": "a12", "templateType": "anything", "description": "", "definition": "random(-1,1,2)"}, "s2": {"group": "Ungrouped variables", "name": "s2", "templateType": "anything", "description": "", "definition": "if(q=m,0.5,-0.5)"}, "ab13": {"group": "Ungrouped variables", "name": "ab13", "templateType": "anything", "description": "", "definition": "a11*b13+a12*b23+a13*b33"}, "b13": {"group": "Ungrouped variables", "name": "b13", "templateType": "anything", "description": "", "definition": "random(-2..2)"}, "m": {"group": "Ungrouped variables", "name": "m", "templateType": "anything", "description": "", "definition": "random(3..4)"}, "w": {"group": "Ungrouped variables", "name": "w", "templateType": "anything", "description": "", "definition": "random(3..5)"}, "q1": {"group": "Ungrouped variables", "name": "q1", "templateType": "anything", "description": "", "definition": "b11*w1+b12*w2+b13*w3"}, "ba12": {"group": "Ungrouped variables", "name": "ba12", "templateType": "anything", "description": "", "definition": "b11*a12+b12*a22+b13*a32"}, "s8": {"group": "Ungrouped variables", "name": "s8", "templateType": "anything", "description": "", "definition": "if(x=r,0.5,-0.5)"}, "mad": {"group": "Ungrouped variables", "name": "mad", "templateType": "anything", "description": "", "definition": "matrix(repeat(repeat(random(-2..9),q),p))"}, "s4": {"group": "Ungrouped variables", "name": "s4", "templateType": "anything", "description": "", "definition": "if(r=2,0.5,-0.5)"}, "z": {"group": "Ungrouped variables", "name": "z", "templateType": "anything", "description": "", "definition": "random(-2,-1,1,2)"}, "ba23": {"group": "Ungrouped variables", "name": "ba23", "templateType": "anything", "description": "", "definition": "b21*a13+b22*a23+b23*a33"}, "s10": {"group": "Ungrouped variables", "name": "s10", "templateType": "anything", "description": "", "definition": "if(u=3,0.5,-0.5)"}, "w1": {"group": "Ungrouped variables", "name": "w1", "templateType": "anything", "description": "", "definition": "random(4..6)"}, "c13": {"group": "Ungrouped variables", "name": "c13", "templateType": "anything", "description": "", "definition": "random(1..2)"}, "ba11": {"group": "Ungrouped variables", "name": "ba11", "templateType": "anything", "description": "", "definition": "b11*a11+b12*a21+b13*a31"}, "s5": {"group": "Ungrouped variables", "name": "s5", "templateType": "anything", "description": "", "definition": "if(m=3,0.5,-0.5)"}, "y": {"group": "Ungrouped variables", "name": "y", "templateType": "anything", "description": "", "definition": "u+random(0,z)"}, "ba31": {"group": "Ungrouped variables", "name": "ba31", "templateType": "anything", "description": "", "definition": "b31*a11+b32*a21+b33*a31"}, "s1": {"group": "Ungrouped variables", "name": "s1", "templateType": "anything", "description": "", "definition": "if(n=p,0.5,-0.5)"}, "ab22": {"group": "Ungrouped variables", "name": "ab22", "templateType": "anything", "description": "", "definition": "a21*b12+a22*b22+a23*b32"}, "a22": {"group": "Ungrouped variables", "name": "a22", "templateType": "anything", "description": "", "definition": "random(-4,-3,-2,-1,1,2,3,5)"}, "a21": {"group": "Ungrouped variables", "name": "a21", "templateType": "anything", "description": "", "definition": "random(-1,0,1)"}, "p3": {"group": "Ungrouped variables", "name": "p3", "templateType": "anything", "description": "", "definition": "a31*v1+a32*v2+a33*v3"}, "ba32": {"group": "Ungrouped variables", "name": "ba32", "templateType": "anything", "description": "", "definition": "b31*a12+b32*a22+b33*a32"}, "ba21": {"group": "Ungrouped variables", "name": "ba21", "templateType": "anything", "description": "", "definition": "b21*a11+b22*a21+b23*a31"}, "q3": {"group": "Ungrouped variables", "name": "q3", "templateType": "anything", "description": "", "definition": "b31*w1+b32*w2+b33*w3"}, "s": {"group": "Ungrouped variables", "name": "s", "templateType": "anything", "description": "", "definition": "random(4..6)"}, "p2": {"group": "Ungrouped variables", "name": "p2", "templateType": "anything", "description": "", "definition": "a21*v1+a22*v2+a23*v3"}, "p": {"group": "Ungrouped variables", "name": "p", "templateType": "anything", "description": "", "definition": "n+random(0,z)"}, "a31": {"group": "Ungrouped variables", "name": "a31", "templateType": "anything", "description": "", "definition": "random(1..3)"}, "s6": {"group": "Ungrouped variables", "name": "s6", "templateType": "anything", "description": "", "definition": "if(r=3,0.5,-0.5)"}, "ba13": {"group": "Ungrouped variables", "name": "ba13", "templateType": "anything", "description": "", "definition": "b11*a13+b12*a23+b13*a33"}, "mae": {"group": "Ungrouped variables", "name": "mae", "templateType": "anything", "description": "", "definition": "matrix(repeat(repeat(random(-2..9),s),r))"}, "ab23": {"group": "Ungrouped variables", "name": "ab23", "templateType": "anything", "description": "", "definition": "a21*b13+a22*b23+a23*b33"}, "s9": {"group": "Ungrouped variables", "name": "s9", "templateType": "anything", "description": "", "definition": "if(w=3,0.5,-0.5)"}, "x": {"group": "Ungrouped variables", "name": "x", "templateType": "anything", "description": "", "definition": "u+random(0,z)"}, "b12": {"group": "Ungrouped variables", "name": "b12", "templateType": "anything", "description": "", "definition": "random(0,1)"}, "ab12": {"group": "Ungrouped variables", "name": "ab12", "templateType": "anything", "description": "", "definition": "a11*b12+a12*b22+a13*b32"}, "w2": {"group": "Ungrouped variables", "name": "w2", "templateType": "anything", "description": "", "definition": "random(-4..4)"}, "a13": {"group": "Ungrouped variables", "name": "a13", "templateType": "anything", "description": "", "definition": "random(1..2)"}, "ab32": {"group": "Ungrouped variables", "name": "ab32", "templateType": "anything", "description": "", "definition": "a31*b12+a32*b22+a33*b32"}, "a23": {"group": "Ungrouped variables", "name": "a23", "templateType": "anything", "description": "", "definition": "random(-2..2)"}, "c12": {"group": "Ungrouped variables", "name": "c12", "templateType": "anything", "description": "", "definition": "random(1..3)"}, "b11": {"group": "Ungrouped variables", "name": "b11", "templateType": "anything", "description": "", "definition": "random(0,1)"}, "q2": {"group": "Ungrouped variables", "name": "q2", "templateType": "anything", "description": "", "definition": "b21*w1+b22*w2+b23*w3"}, "n": {"group": "Ungrouped variables", "name": "n", "templateType": "anything", "description": "", "definition": "random(3..6 except m)"}, "a11": {"group": "Ungrouped variables", "name": "a11", "templateType": "anything", "description": "", "definition": "random(1..4)"}, "a33": {"group": "Ungrouped variables", "name": "a33", "templateType": "anything", "description": "", "definition": "random(1..4)"}, "v2": {"group": "Ungrouped variables", "name": "v2", "templateType": "anything", "description": "", "definition": "random(1..4)"}, "p1": {"group": "Ungrouped variables", "name": "p1", "templateType": "anything", "description": "", "definition": "a11*v1+a12*v2+a13*v3"}, "v": {"group": "Ungrouped variables", "name": "v", "templateType": "anything", "description": "", "definition": "[[s1,-s1],[s2,-s2],[s3,-s3],[s4,-s4],[s5,-s5],[s6,-s6],[s7,-s7],[s8,-s8],[s9,-s9],[s10,-s10]]"}, "b22": {"group": "Ungrouped variables", "name": "b22", "templateType": "anything", "description": "", "definition": "random(-4,-3,-2,-1,1,2,3,4)"}}, "tags": [], "preamble": {"js": "", "css": ""}, "parts": [{"adaptiveMarkingPenalty": 0, "marks": 0, "customMarkingAlgorithm": "", "showFeedbackIcon": true, "variableReplacements": [], "type": "gapfill", "extendBaseMarkingAlgorithm": true, "unitTests": [], "gaps": [{"adaptiveMarkingPenalty": 0, "marks": 1, "correctAnswer": "matrix([[p1],[p2],[p3]])", "customMarkingAlgorithm": "", "showFeedbackIcon": true, "variableReplacements": [], "numRows": "3", "type": "matrix", "correctAnswerFractions": false, "variableReplacementStrategy": "originalfirst", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "allowResize": false, "useCustomName": false, "numColumns": 1, "scripts": {}, "tolerance": 0, "allowFractions": false, "markPerCell": false, "customName": ""}, {"adaptiveMarkingPenalty": 0, "marks": 1, "correctAnswer": "matrix([[q1],[q2],[q3]])", "customMarkingAlgorithm": "", "showFeedbackIcon": true, "variableReplacements": [], "numRows": "3", "type": "matrix", "correctAnswerFractions": false, "variableReplacementStrategy": "originalfirst", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "allowResize": false, "useCustomName": false, "numColumns": 1, "scripts": {}, "tolerance": 0, "allowFractions": false, "markPerCell": false, "customName": ""}, {"adaptiveMarkingPenalty": 0, "marks": 1, "correctAnswer": "matrix([\n    [ba11,ba12,ba13],\n    [ba21,ba22,ba23],\n    [ba31,ba32,ba33]\n])", "customMarkingAlgorithm": "", "showFeedbackIcon": true, "variableReplacements": [], "numRows": "3", "type": "matrix", "correctAnswerFractions": false, "variableReplacementStrategy": "originalfirst", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "allowResize": false, "useCustomName": false, "numColumns": "3", "scripts": {}, "tolerance": 0, "allowFractions": false, "markPerCell": false, "customName": ""}, {"adaptiveMarkingPenalty": 0, "marks": 1, "correctAnswer": "matrix([\n    [ab11,ab12,ab13],\n    [ab21,ab22,ab23],\n    [ab31,ab32,ab33]\n])", "customMarkingAlgorithm": "", "showFeedbackIcon": true, "variableReplacements": [], "numRows": "3", "type": "matrix", "correctAnswerFractions": false, "variableReplacementStrategy": "originalfirst", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "allowResize": false, "useCustomName": false, "numColumns": "3", "scripts": {}, "tolerance": 0, "allowFractions": false, "markPerCell": false, "customName": ""}], "showCorrectAnswer": true, "customName": "", "sortAnswers": false, "useCustomName": false, "scripts": {}, "prompt": "<p>Let</p>\n<p>\\[A = \\left(\\begin{array}{rrr} \\var{a11} &amp; \\var{a12} &amp; \\var{a13}\\\\ \\var{a21} &amp; \\var{a22} &amp; \\var{a23}\\\\ \\var{a31} &amp; \\var{a32} &amp; \\var{a33}\\\\ \\end{array}\\right),\\;\\;\\;\\; B= \\left(\\begin{array}{rrr} \\var{b11} &amp; \\var{b12} &amp; \\var{b13}\\\\ \\var{b21} &amp; \\var{b22} &amp; \\var{b23}\\\\ \\var{b31} &amp; \\var{b32} &amp; \\var{b33}\\\\ \\end{array}\\right),\\;\\;\\;\\; v= \\left(\\begin{array}{r} \\var{v1}\\\\ \\var{v2} \\\\ \\var{v3} \\end{array}\\right),\\;\\;\\;\\; w= \\left(\\begin{array}{r} \\var{w1}\\\\ \\var{w2} \\\\ \\var{w3} \\end{array}\\right)\\]</p>\n<p>Find the following products:</p>\n<p>$Av = $ [[0]]</p>\n<p>$Bw = $ [[1]]</p>\n<p>$BA = $ [[2]]</p>\n<p>$AB = $ [[3]]</p>", "variableReplacementStrategy": "originalfirst"}, {"adaptiveMarkingPenalty": 0, "marks": 0, "customMarkingAlgorithm": "", "showFeedbackIcon": true, "variableReplacements": [], "type": "gapfill", "extendBaseMarkingAlgorithm": true, "unitTests": [], "gaps": [{"showFeedbackIcon": true, "type": "m_n_x", "showCellAnswerState": true, "extendBaseMarkingAlgorithm": true, "maxMarks": 0, "minMarks": 0, "displayType": "radiogroup", "customName": "", "shuffleChoices": true, "scripts": {}, "customMarkingAlgorithm": "", "matrix": "v", "adaptiveMarkingPenalty": 0, "marks": 0, "choices": ["<p>$CD$</p>", "<p>$DC$</p>", "<p>$EF$</p>", "<p>$FE$</p>", "<p>$BC$</p>", "<p>$AE$</p>", "<p>$GH$</p>", "<p>$HE$</p>", "<p>$AG$</p>", "<p>$GB$</p>"], "variableReplacementStrategy": "originalfirst", "minAnswers": 0, "variableReplacements": [], "unitTests": [], "warningType": "none", "maxAnswers": 0, "useCustomName": false, "showCorrectAnswer": true, "shuffleAnswers": false, "answers": ["<strong>Can be calculated</strong>", "<strong>Cannot be calculated</strong>"], "layout": {"expression": ""}}], "showCorrectAnswer": true, "customName": "", "sortAnswers": false, "useCustomName": false, "scripts": {}, "prompt": "<p>Consider the following matrices together with the matrices from the first part of the question.</p>\n<p>\\[\\begin{eqnarray}&amp;C=&amp; \\var{mac},\\;\\;\\;\\; &amp;D=&amp; \\var{mad},\\;\\;\\; \\;&amp;E= &amp;\\var{mae}\\\\&amp;F=&amp; \\left(\\begin{array}{rr} \\var{w1} &amp; \\var{a12}\\\\ \\var{w2} &amp; \\var{b23} \\\\ \\var{w3} &amp; \\var{w2} \\\\\\var{v1} &amp; \\var{b12}\\\\ 0 &amp; \\var{-w2} \\end{array}\\right),\\;\\;\\;\\;&amp;G=&amp;\\var{mag},\\;\\;\\;\\;&amp;H=&amp;\\var{mah} \\end{eqnarray}\\]</p>\n<p>Which of the following products of matrices can be calculated?</p>\n<p>[[0]]</p>\n<p><strong>Please note that for every correct answer you get 0.5 marks and for every incorrect answer 0.5 is taken away. The minimum mark you can get is 0.</strong></p>", "variableReplacementStrategy": "originalfirst"}], "statement": "\n  <p>Answer the following questions on matrices.</p>\n  <p>&nbsp;</p>\n  "}, {"name": "L1 - Matrix multiplication (2x2)", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false}, "contributors": [{"name": "Newcastle University Mathematics and Statistics", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/697/"}, {"name": "Timur Zaripov", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/3272/"}], "functions": {}, "rulesets": {"std": ["all", "fractionNumbers", "!collectNumbers"]}, "metadata": {"licence": "Creative Commons Attribution 4.0 International", "description": "<p>Multiplication of $2 \\times 2$ matrices.</p>"}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["a", "b", "c"], "variable_groups": [], "preamble": {"js": "", "css": ""}, "variables": {"c": {"definition": "matrix([ [random(1,0,4), a[0][1]+b[0][1]], [random(2..5),random(0,1)] ])", "name": "c", "templateType": "anything", "description": "", "group": "Ungrouped variables"}, "a": {"definition": "matrix([ [random(-2,1,2),random(1..4)], [random(-2..2),random(1..3)] ])", "name": "a", "templateType": "anything", "description": "", "group": "Ungrouped variables"}, "b": {"definition": "matrix([ [random(-3,-1,0,3),random(-3..1)], [random(2,3),random(-3..-1)] ])", "name": "b", "templateType": "anything", "description": "", "group": "Ungrouped variables"}}, "tags": [], "advice": "<h4>a)</h4>\n<p>\\begin{align}<br/>\\mathbf{AB} &amp;= \\var{A}\\var{B} \\\\<br/>&amp;= \\begin{pmatrix} \\simplify[]{ {a[0][0]}*{b[0][0]}+{a[0][1]}*{b[1][0]} } &amp; \\simplify[]{ {a[0][0]}*{b[0][1]} + {a[0][1]}*{b[1][1]} } \\\\ \\simplify[]{ {a[1][0]}*{b[0][0]} + {a[1][1]}*{b[1][0]} } &amp; \\simplify[]{ {a[1][0]}*{b[0][1]} + {a[1][1]}*{b[1][1]} } \\end{pmatrix} \\\\<br/>&amp;= \\var{a*b}<br/>\\end{align}</p>\n<h4>b)</h4>\n<p>\\begin{align}<br/>\\mathbf{BA} &amp;= \\var{B}\\var{A} \\\\<br/>&amp;= \\begin{pmatrix} \\simplify[]{ {b[0][0]}*{a[0][0]}+{b[0][1]}*{a[1][0]} } &amp; \\simplify[]{ {b[0][0]}*{a[0][1]} + {b[0][1]}*{a[1][1]} } \\\\ \\simplify[]{ {b[1][0]}*{a[0][0]} + {b[1][1]}*{a[1][0]} } &amp; \\simplify[]{ {b[1][0]}*{a[0][1]} + {b[1][1]}*{a[1][1]} } \\end{pmatrix} \\\\<br/>&amp;= \\var{b*a}<br/>\\end{align}</p>\n<h4>c)</h4>\n<p>\\begin{align}<br/>\\mathbf{CB} &amp;= \\var{C}\\var{B} \\\\<br/>&amp;= \\begin{pmatrix} \\simplify[]{ {c[0][0]}*{b[0][0]}+{c[0][1]}*{b[1][0]} } &amp; \\simplify[]{ {c[0][0]}*{b[0][1]} + {c[0][1]}*{b[1][1]} } \\\\ \\simplify[]{ {c[1][0]}*{b[0][0]} + {c[1][1]}*{b[1][0]} } &amp; \\simplify[]{ {c[1][0]}*{b[0][1]} + {c[1][1]}*{b[1][1]} } \\end{pmatrix} \\\\<br/>&amp;= \\var{c*b}<br/>\\end{align}</p>\n<h4>d)</h4>\n<p>\\begin{align}<br/>\\mathbf{AC} &amp;= \\var{A}\\var{C} \\\\<br/>&amp;= \\begin{pmatrix} \\simplify[]{ {a[0][0]}*{c[0][0]}+{a[0][1]}*{c[1][0]} } &amp; \\simplify[]{ {a[0][0]}*{c[0][1]} + {a[0][1]}*{c[1][1]} } \\\\ \\simplify[]{ {a[1][0]}*{c[0][0]} + {a[1][1]}*{c[1][0]} } &amp; \\simplify[]{ {a[1][0]}*{c[0][1]} + {a[1][1]}*{c[1][1]} } \\end{pmatrix} \\\\<br/>&amp;= \\var{a*c}<br/>\\end{align}</p>", "parts": [{"adaptiveMarkingPenalty": 0, "marks": 0, "customMarkingAlgorithm": "", "showFeedbackIcon": true, "variableReplacements": [], "type": "gapfill", "extendBaseMarkingAlgorithm": true, "unitTests": [], "gaps": [{"adaptiveMarkingPenalty": 0, "marks": 1, "customMarkingAlgorithm": "", "showFeedbackIcon": true, "variableReplacements": [], "numRows": "2", "type": "matrix", "correctAnswerFractions": false, "extendBaseMarkingAlgorithm": true, "numColumns": "2", "unitTests": [], "showCorrectAnswer": true, "allowResize": false, "useCustomName": false, "customName": "", "correctAnswer": "a*b", "scripts": {}, "tolerance": 0, "allowFractions": false, "markPerCell": false, "variableReplacementStrategy": "originalfirst"}], "showCorrectAnswer": true, "customName": "", "sortAnswers": false, "useCustomName": false, "scripts": {}, "prompt": "<p style=\"text-align: center;\">$\\mathbf{AB} = \\var{A}\\var{B} = $&nbsp;[[0]]</p>", "variableReplacementStrategy": "originalfirst"}, {"adaptiveMarkingPenalty": 0, "marks": 0, "customMarkingAlgorithm": "", "showFeedbackIcon": true, "variableReplacements": [], "type": "gapfill", "extendBaseMarkingAlgorithm": true, "unitTests": [], "gaps": [{"adaptiveMarkingPenalty": 0, "marks": 1, "customMarkingAlgorithm": "", "showFeedbackIcon": true, "variableReplacements": [], "numRows": "2", "type": "matrix", "correctAnswerFractions": false, "extendBaseMarkingAlgorithm": true, "numColumns": "2", "unitTests": [], "showCorrectAnswer": true, "allowResize": false, "useCustomName": false, "customName": "", "correctAnswer": "b*a", "scripts": {}, "tolerance": 0, "allowFractions": false, "markPerCell": false, "variableReplacementStrategy": "originalfirst"}], "showCorrectAnswer": true, "customName": "", "sortAnswers": false, "useCustomName": false, "scripts": {}, "prompt": "<p style=\"text-align: center;\">$\\mathbf{BA} = \\var{B}\\var{A} = $&nbsp;[[0]]</p>", "variableReplacementStrategy": "originalfirst"}, {"adaptiveMarkingPenalty": 0, "marks": 0, "customMarkingAlgorithm": "", "showFeedbackIcon": true, "variableReplacements": [], "type": "gapfill", "extendBaseMarkingAlgorithm": true, "unitTests": [], "gaps": [{"adaptiveMarkingPenalty": 0, "marks": 1, "customMarkingAlgorithm": "", "showFeedbackIcon": true, "variableReplacements": [], "numRows": "2", "type": "matrix", "correctAnswerFractions": false, "extendBaseMarkingAlgorithm": true, "numColumns": "2", "unitTests": [], "showCorrectAnswer": true, "allowResize": false, "useCustomName": false, "customName": "", "correctAnswer": "c*b", "scripts": {}, "tolerance": 0, "allowFractions": false, "markPerCell": false, "variableReplacementStrategy": "originalfirst"}], "showCorrectAnswer": true, "customName": "", "sortAnswers": false, "useCustomName": false, "scripts": {}, "prompt": "<p style=\"text-align: center;\">$\\mathbf{CB} = \\var{C}\\var{B} = $&nbsp;[[0]]</p>", "variableReplacementStrategy": "originalfirst"}, {"adaptiveMarkingPenalty": 0, "marks": 0, "customMarkingAlgorithm": "", "showFeedbackIcon": true, "variableReplacements": [], "type": "gapfill", "extendBaseMarkingAlgorithm": true, "unitTests": [], "gaps": [{"adaptiveMarkingPenalty": 0, "marks": 1, "customMarkingAlgorithm": "", "showFeedbackIcon": true, "variableReplacements": [], "numRows": "2", "type": "matrix", "correctAnswerFractions": false, "extendBaseMarkingAlgorithm": true, "numColumns": "2", "unitTests": [], "showCorrectAnswer": true, "allowResize": false, "useCustomName": false, "customName": "", "correctAnswer": "a*c", "scripts": {}, "tolerance": 0, "allowFractions": false, "markPerCell": false, "variableReplacementStrategy": "originalfirst"}], "showCorrectAnswer": true, "customName": "", "sortAnswers": false, "useCustomName": false, "scripts": {}, "prompt": "<p style=\"text-align: center;\">$\\mathbf{AC} = \\var{A}\\var{C} = $&nbsp;[[0]]</p>", "variableReplacementStrategy": "originalfirst"}], "statement": "<p>Let</p>\n<p>\\begin{align} \\mathbf{A} &amp;= \\var{a}, &amp; \\mathbf{B} &amp;= \\var{b}, &amp; \\mathbf{C} &amp;= \\var{c} \\end{align}</p>\n<p>Calculate the following products of these matrices:</p>"}, {"name": "Matrix Arithmetics", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false}, "contributors": [{"name": "Marie Nicholson", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/1799/"}, {"name": "Timur Zaripov", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/3272/"}], "tags": [], "metadata": {"description": "<p>Addition, subtraction and multiplication by a scalar for 2 x 2 matrices.</p>", "licence": "Creative Commons Attribution 4.0 International"}, "statement": "<p>Find the answers to the following matrix calculations, filling in your answers in the spaces provided.</p>", "advice": "<p>(a)</p>\n<p>$\\var{a}+\\var{b}=\\var{u}$</p>\n<p><strong>Just add corresponding elements together.</strong></p>\n<p>(b)</p>\n<p>$\\var{p}\\var{c}-\\var{q}\\var{d}$</p>\n<p><strong>Multiply each element of the first matrix by $\\var{p}$...</strong></p>\n<p>$\\var{p}\\var{c}=\\var{pc}$</p>\n<p><strong>...and multiply each element of the second matrix by $\\var{q}$.<br/></strong></p>\n<p>$\\var{q}\\var{d}=\\var{qd}$</p>\n<p><strong>Finally, subtract corresponding elements</strong></p>\n<p>$\\var{p}\\var{c}-\\var{q}\\var{d}=\\var{pc}-\\var{qd}=\\var{v}$</p>", "rulesets": {"ruleset0": []}, "variables": {"at": {"name": "at", "group": "Ungrouped variables", "definition": "transpose(a)", "description": "", "templateType": "anything"}, "a12": {"name": "a12", "group": "Ungrouped variables", "definition": "random(-10..10#1)", "description": "", "templateType": "anything"}, "ans": {"name": "ans", "group": "Ungrouped variables", "definition": "a+b", "description": "", "templateType": "anything"}, "a11": {"name": "a11", "group": "Ungrouped variables", "definition": "random(-10..10#1)", "description": "", "templateType": "anything"}, "c21": {"name": "c21", "group": "Ungrouped variables", "definition": "random(-10..10#1)", "description": "", "templateType": "anything"}, "bat": {"name": "bat", "group": "Ungrouped variables", "definition": "b*at", "description": "", "templateType": "anything"}, "b22": {"name": "b22", "group": "Ungrouped variables", "definition": "random(-10..10#1)", "description": "", "templateType": "anything"}, "p": {"name": "p", "group": "Ungrouped variables", "definition": "random(-10..10#1 except 0 except 1 except q)", "description": "", "templateType": "anything"}, "deta": {"name": "deta", "group": "Ungrouped variables", "definition": "det(a)", "description": "", "templateType": "anything"}, "v": {"name": "v", "group": "Ungrouped variables", "definition": "pc-qd", "description": "", "templateType": "anything"}, "a22": {"name": "a22", "group": "Ungrouped variables", "definition": "if(test*a11=a21*a12,test+1,test)", "description": "", "templateType": "anything"}, "a21": {"name": "a21", "group": "Ungrouped variables", "definition": "random(-10..10#1)", "description": "", "templateType": "anything"}, "u": {"name": "u", "group": "Ungrouped variables", "definition": "a+b", "description": "", "templateType": "anything"}, "d21": {"name": "d21", "group": "Ungrouped variables", "definition": "random(-10..10#1)", "description": "", "templateType": "anything"}, "test": {"name": "test", "group": "Ungrouped variables", "definition": "random(-10..10#1)", "description": "", "templateType": "anything"}, "c12": {"name": "c12", "group": "Ungrouped variables", "definition": "random(-10..10#1)", "description": "", "templateType": "anything"}, "b": {"name": "b", "group": "Ungrouped variables", "definition": "matrix([b11,b12],[b21,b22])", "description": "", "templateType": "anything"}, "q": {"name": "q", "group": "Ungrouped variables", "definition": "random(2..10#1)", "description": "", "templateType": "anything"}, "d": {"name": "d", "group": "Ungrouped variables", "definition": "matrix([d11,d12],[d21,d22])", "description": "", "templateType": "anything"}, "det": {"name": "det", "group": "Ungrouped variables", "definition": "a11*a22-a21*a12", "description": "", "templateType": "anything"}, "b11": {"name": "b11", "group": "Ungrouped variables", "definition": "random(-10..10#1)", "description": "", "templateType": "anything"}, "qd": {"name": "qd", "group": "Ungrouped variables", "definition": "q*d", "description": "", "templateType": "anything"}, "d12": {"name": "d12", "group": "Ungrouped variables", "definition": "random(-10..10#1)", "description": "", "templateType": "anything"}, "b21": {"name": "b21", "group": "Ungrouped variables", "definition": "random(-10..10#1)", "description": "", "templateType": "anything"}, "a": {"name": "a", "group": "Ungrouped variables", "definition": "matrix([a11,a12],[a21,a22])", "description": "", "templateType": "anything"}, "c11": {"name": "c11", "group": "Ungrouped variables", "definition": "random(-10..10#1)", "description": "", "templateType": "anything"}, "b12": {"name": "b12", "group": "Ungrouped variables", "definition": "random(-10..10#1)", "description": "", "templateType": "anything"}, "adj": {"name": "adj", "group": "Ungrouped variables", "definition": "matrix([a22,-a12],[-a21,a11])", "description": "", "templateType": "anything"}, "d11": {"name": "d11", "group": "Ungrouped variables", "definition": "random(-10..10#1)", "description": "", "templateType": "anything"}, "c": {"name": "c", "group": "Ungrouped variables", "definition": "matrix([c11,c12],[c21,c22])", "description": "", "templateType": "anything"}, "c22": {"name": "c22", "group": "Ungrouped variables", "definition": "random(-10..10#1)", "description": "", "templateType": "anything"}, "d22": {"name": "d22", "group": "Ungrouped variables", "definition": "random(-10..10#1)", "description": "", "templateType": "anything"}, "ans2": {"name": "ans2", "group": "Ungrouped variables", "definition": "c-d", "description": "", "templateType": "anything"}, "pc": {"name": "pc", "group": "Ungrouped variables", "definition": "p*c", "description": "", "templateType": "anything"}, "ab": {"name": "ab", "group": "Ungrouped variables", "definition": "a*b", "description": "", "templateType": "anything"}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["a21", "a22", "ab", "ans2", "b22", "b21", "at", "ans", "d11", "d12", "p", "pc", "b12", "b11", "test", "c12", "c11", "adj", "a", "a11", "a12", "c", "b", "c22", "bat", "c21", "d", "deta", "d22", "det", "q", "d21", "u", "qd", "v"], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p class=\"MTDisplayEquation\">$\\var{a}+\\var{b}=$<span style=\"mso-tab-count: 1;\">[[0]]&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp; </span></p>\n<p></p>\n<p><!-- [if gte mso 9]><xml>\n <o:OfficeDocumentSettings>\n  <o:AllowPNG/>\n </o:OfficeDocumentSettings>\n</xml><![endif]--></p>\n<p><!-- [if gte mso 9]><xml>\n <w:WordDocument>\n  <w:View>Normal</w:View>\n  <w:Zoom>0</w:Zoom>\n  <w:TrackMoves/>\n  <w:TrackFormatting/>\n  <w:PunctuationKerning/>\n  <w:ValidateAgainstSchemas/>\n  <w:SaveIfXMLInvalid>false</w:SaveIfXMLInvalid>\n  <w:IgnoreMixedContent>false</w:IgnoreMixedContent>\n  <w:AlwaysShowPlaceholderText>false</w:AlwaysShowPlaceholderText>\n  <w:DoNotPromoteQF/>\n  <w:LidThemeOther>EN-GB</w:LidThemeOther>\n  <w:LidThemeAsian>X-NONE</w:LidThemeAsian>\n  <w:LidThemeComplexScript>X-NONE</w:LidThemeComplexScript>\n  <w:Compatibility>\n   <w:BreakWrappedTables/>\n   <w:SnapToGridInCell/>\n   <w:WrapTextWithPunct/>\n   <w:UseAsianBreakRules/>\n   <w:DontGrowAutofit/>\n   <w:SplitPgBreakAndParaMark/>\n   <w:EnableOpenTypeKerning/>\n   <w:DontFlipMirrorIndents/>\n   <w:OverrideTableStyleHps/>\n  </w:Compatibility>\n  <m:mathPr>\n   <m:mathFont m:val=\"Cambria Math\"/>\n   <m:brkBin m:val=\"before\"/>\n   <m:brkBinSub m:val=\"&#45;-\"/>\n   <m:smallFrac m:val=\"off\"/>\n   <m:dispDef/>\n   <m:lMargin m:val=\"0\"/>\n   <m:rMargin m:val=\"0\"/>\n   <m:defJc m:val=\"centerGroup\"/>\n   <m:wrapIndent m:val=\"1440\"/>\n   <m:intLim m:val=\"subSup\"/>\n   <m:naryLim m:val=\"undOvr\"/>\n  </m:mathPr></w:WordDocument>\n</xml><![endif]--><!-- [if gte mso 9]><xml>\n <w:LatentStyles DefLockedState=\"false\" DefUnhideWhenUsed=\"true\"\n  DefSemiHidden=\"true\" DefQFormat=\"false\" DefPriority=\"99\"\n  LatentStyleCount=\"267\">\n  <w:LsdException Locked=\"false\" Priority=\"0\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" QFormat=\"true\" Name=\"Normal\"/>\n  <w:LsdException Locked=\"false\" Priority=\"9\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" QFormat=\"true\" Name=\"heading 1\"/>\n  <w:LsdException Locked=\"false\" Priority=\"9\" QFormat=\"true\" Name=\"heading 2\"/>\n  <w:LsdException Locked=\"false\" Priority=\"9\" QFormat=\"true\" Name=\"heading 3\"/>\n  <w:LsdException Locked=\"false\" Priority=\"9\" QFormat=\"true\" Name=\"heading 4\"/>\n  <w:LsdException Locked=\"false\" Priority=\"9\" QFormat=\"true\" Name=\"heading 5\"/>\n  <w:LsdException Locked=\"false\" Priority=\"9\" QFormat=\"true\" Name=\"heading 6\"/>\n  <w:LsdException Locked=\"false\" Priority=\"9\" QFormat=\"true\" Name=\"heading 7\"/>\n  <w:LsdException Locked=\"false\" Priority=\"9\" QFormat=\"true\" Name=\"heading 8\"/>\n  <w:LsdException Locked=\"false\" Priority=\"9\" QFormat=\"true\" Name=\"heading 9\"/>\n  <w:LsdException Locked=\"false\" Priority=\"39\" Name=\"toc 1\"/>\n  <w:LsdException Locked=\"false\" Priority=\"39\" Name=\"toc 2\"/>\n  <w:LsdException Locked=\"false\" Priority=\"39\" Name=\"toc 3\"/>\n  <w:LsdException Locked=\"false\" Priority=\"39\" Name=\"toc 4\"/>\n  <w:LsdException Locked=\"false\" Priority=\"39\" Name=\"toc 5\"/>\n  <w:LsdException Locked=\"false\" Priority=\"39\" Name=\"toc 6\"/>\n  <w:LsdException Locked=\"false\" Priority=\"39\" Name=\"toc 7\"/>\n  <w:LsdException Locked=\"false\" Priority=\"39\" Name=\"toc 8\"/>\n  <w:LsdException Locked=\"false\" Priority=\"39\" Name=\"toc 9\"/>\n  <w:LsdException Locked=\"false\" Priority=\"35\" QFormat=\"true\" Name=\"caption\"/>\n  <w:LsdException Locked=\"false\" Priority=\"10\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" QFormat=\"true\" Name=\"Title\"/>\n  <w:LsdException Locked=\"false\" Priority=\"1\" Name=\"Default Paragraph Font\"/>\n  <w:LsdException Locked=\"false\" Priority=\"11\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" QFormat=\"true\" Name=\"Subtitle\"/>\n  <w:LsdException Locked=\"false\" Priority=\"22\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" QFormat=\"true\" Name=\"Strong\"/>\n  <w:LsdException Locked=\"false\" Priority=\"20\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" QFormat=\"true\" Name=\"Emphasis\"/>\n  <w:LsdException Locked=\"false\" Priority=\"59\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Table Grid\"/>\n  <w:LsdException Locked=\"false\" UnhideWhenUsed=\"false\" Name=\"Placeholder Text\"/>\n  <w:LsdException Locked=\"false\" Priority=\"1\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" QFormat=\"true\" Name=\"No Spacing\"/>\n  <w:LsdException Locked=\"false\" Priority=\"60\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Light Shading\"/>\n  <w:LsdException Locked=\"false\" Priority=\"61\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Light List\"/>\n  <w:LsdException Locked=\"false\" Priority=\"62\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Light Grid\"/>\n  <w:LsdException Locked=\"false\" Priority=\"63\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Shading 1\"/>\n  <w:LsdException Locked=\"false\" Priority=\"64\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Shading 2\"/>\n  <w:LsdException Locked=\"false\" Priority=\"65\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium List 1\"/>\n  <w:LsdException Locked=\"false\" Priority=\"66\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium List 2\"/>\n  <w:LsdException Locked=\"false\" Priority=\"67\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Grid 1\"/>\n  <w:LsdException Locked=\"false\" Priority=\"68\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Grid 2\"/>\n  <w:LsdException Locked=\"false\" Priority=\"69\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Grid 3\"/>\n  <w:LsdException Locked=\"false\" Priority=\"70\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Dark List\"/>\n  <w:LsdException Locked=\"false\" Priority=\"71\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Colorful Shading\"/>\n  <w:LsdException Locked=\"false\" Priority=\"72\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Colorful List\"/>\n  <w:LsdException Locked=\"false\" Priority=\"73\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Colorful Grid\"/>\n  <w:LsdException Locked=\"false\" Priority=\"60\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Light Shading Accent 1\"/>\n  <w:LsdException Locked=\"false\" Priority=\"61\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Light List Accent 1\"/>\n  <w:LsdException Locked=\"false\" Priority=\"62\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Light Grid Accent 1\"/>\n  <w:LsdException Locked=\"false\" Priority=\"63\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Shading 1 Accent 1\"/>\n  <w:LsdException Locked=\"false\" Priority=\"64\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Shading 2 Accent 1\"/>\n  <w:LsdException Locked=\"false\" Priority=\"65\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium List 1 Accent 1\"/>\n  <w:LsdException Locked=\"false\" UnhideWhenUsed=\"false\" Name=\"Revision\"/>\n  <w:LsdException Locked=\"false\" Priority=\"34\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" QFormat=\"true\" Name=\"List Paragraph\"/>\n  <w:LsdException Locked=\"false\" Priority=\"29\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" QFormat=\"true\" Name=\"Quote\"/>\n  <w:LsdException Locked=\"false\" Priority=\"30\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" QFormat=\"true\" Name=\"Intense Quote\"/>\n  <w:LsdException Locked=\"false\" Priority=\"66\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium List 2 Accent 1\"/>\n  <w:LsdException Locked=\"false\" Priority=\"67\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Grid 1 Accent 1\"/>\n  <w:LsdException Locked=\"false\" Priority=\"68\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Grid 2 Accent 1\"/>\n  <w:LsdException Locked=\"false\" Priority=\"69\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Grid 3 Accent 1\"/>\n  <w:LsdException Locked=\"false\" Priority=\"70\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Dark List Accent 1\"/>\n  <w:LsdException Locked=\"false\" Priority=\"71\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Colorful Shading Accent 1\"/>\n  <w:LsdException Locked=\"false\" Priority=\"72\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Colorful List Accent 1\"/>\n  <w:LsdException Locked=\"false\" Priority=\"73\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Colorful Grid Accent 1\"/>\n  <w:LsdException Locked=\"false\" Priority=\"60\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Light Shading Accent 2\"/>\n  <w:LsdException Locked=\"false\" Priority=\"61\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Light List Accent 2\"/>\n  <w:LsdException Locked=\"false\" Priority=\"62\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Light Grid Accent 2\"/>\n  <w:LsdException Locked=\"false\" Priority=\"63\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Shading 1 Accent 2\"/>\n  <w:LsdException Locked=\"false\" Priority=\"64\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Shading 2 Accent 2\"/>\n  <w:LsdException Locked=\"false\" Priority=\"65\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium List 1 Accent 2\"/>\n  <w:LsdException Locked=\"false\" Priority=\"66\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium List 2 Accent 2\"/>\n  <w:LsdException Locked=\"false\" Priority=\"67\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Grid 1 Accent 2\"/>\n  <w:LsdException Locked=\"false\" Priority=\"68\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Grid 2 Accent 2\"/>\n  <w:LsdException Locked=\"false\" Priority=\"69\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Grid 3 Accent 2\"/>\n  <w:LsdException Locked=\"false\" Priority=\"70\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Dark List Accent 2\"/>\n  <w:LsdException Locked=\"false\" Priority=\"71\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Colorful Shading Accent 2\"/>\n  <w:LsdException Locked=\"false\" Priority=\"72\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Colorful List Accent 2\"/>\n  <w:LsdException Locked=\"false\" Priority=\"73\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Colorful Grid Accent 2\"/>\n  <w:LsdException Locked=\"false\" Priority=\"60\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Light Shading Accent 3\"/>\n  <w:LsdException Locked=\"false\" Priority=\"61\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Light List Accent 3\"/>\n  <w:LsdException Locked=\"false\" Priority=\"62\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Light Grid Accent 3\"/>\n  <w:LsdException Locked=\"false\" Priority=\"63\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Shading 1 Accent 3\"/>\n  <w:LsdException Locked=\"false\" Priority=\"64\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Shading 2 Accent 3\"/>\n  <w:LsdException Locked=\"false\" Priority=\"65\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium List 1 Accent 3\"/>\n  <w:LsdException Locked=\"false\" Priority=\"66\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium List 2 Accent 3\"/>\n  <w:LsdException Locked=\"false\" Priority=\"67\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Grid 1 Accent 3\"/>\n  <w:LsdException Locked=\"false\" Priority=\"68\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Grid 2 Accent 3\"/>\n  <w:LsdException Locked=\"false\" Priority=\"69\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Grid 3 Accent 3\"/>\n  <w:LsdException Locked=\"false\" Priority=\"70\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Dark List Accent 3\"/>\n  <w:LsdException Locked=\"false\" Priority=\"71\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Colorful Shading Accent 3\"/>\n  <w:LsdException Locked=\"false\" Priority=\"72\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Colorful List Accent 3\"/>\n  <w:LsdException Locked=\"false\" Priority=\"73\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Colorful Grid Accent 3\"/>\n  <w:LsdException Locked=\"false\" Priority=\"60\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Light Shading Accent 4\"/>\n  <w:LsdException Locked=\"false\" Priority=\"61\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Light List Accent 4\"/>\n  <w:LsdException Locked=\"false\" Priority=\"62\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Light Grid Accent 4\"/>\n  <w:LsdException Locked=\"false\" Priority=\"63\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Shading 1 Accent 4\"/>\n  <w:LsdException Locked=\"false\" Priority=\"64\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Shading 2 Accent 4\"/>\n  <w:LsdException Locked=\"false\" Priority=\"65\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium List 1 Accent 4\"/>\n  <w:LsdException Locked=\"false\" Priority=\"66\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium List 2 Accent 4\"/>\n  <w:LsdException Locked=\"false\" Priority=\"67\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Grid 1 Accent 4\"/>\n  <w:LsdException Locked=\"false\" Priority=\"68\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Grid 2 Accent 4\"/>\n  <w:LsdException Locked=\"false\" Priority=\"69\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Grid 3 Accent 4\"/>\n  <w:LsdException Locked=\"false\" Priority=\"70\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Dark List Accent 4\"/>\n  <w:LsdException Locked=\"false\" Priority=\"71\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Colorful Shading Accent 4\"/>\n  <w:LsdException Locked=\"false\" Priority=\"72\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Colorful List Accent 4\"/>\n  <w:LsdException Locked=\"false\" Priority=\"73\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Colorful Grid Accent 4\"/>\n  <w:LsdException Locked=\"false\" Priority=\"60\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Light Shading Accent 5\"/>\n  <w:LsdException Locked=\"false\" Priority=\"61\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Light List Accent 5\"/>\n  <w:LsdException Locked=\"false\" Priority=\"62\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Light Grid Accent 5\"/>\n  <w:LsdException Locked=\"false\" Priority=\"63\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Shading 1 Accent 5\"/>\n  <w:LsdException Locked=\"false\" Priority=\"64\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Shading 2 Accent 5\"/>\n  <w:LsdException Locked=\"false\" Priority=\"65\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium List 1 Accent 5\"/>\n  <w:LsdException Locked=\"false\" Priority=\"66\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium List 2 Accent 5\"/>\n  <w:LsdException Locked=\"false\" Priority=\"67\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Grid 1 Accent 5\"/>\n  <w:LsdException Locked=\"false\" Priority=\"68\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Grid 2 Accent 5\"/>\n  <w:LsdException Locked=\"false\" Priority=\"69\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Grid 3 Accent 5\"/>\n  <w:LsdException Locked=\"false\" Priority=\"70\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Dark List Accent 5\"/>\n  <w:LsdException Locked=\"false\" Priority=\"71\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Colorful Shading Accent 5\"/>\n  <w:LsdException Locked=\"false\" Priority=\"72\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Colorful List Accent 5\"/>\n  <w:LsdException Locked=\"false\" Priority=\"73\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Colorful Grid Accent 5\"/>\n  <w:LsdException Locked=\"false\" Priority=\"60\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Light Shading Accent 6\"/>\n  <w:LsdException Locked=\"false\" Priority=\"61\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Light List Accent 6\"/>\n  <w:LsdException Locked=\"false\" Priority=\"62\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Light Grid Accent 6\"/>\n  <w:LsdException Locked=\"false\" Priority=\"63\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Shading 1 Accent 6\"/>\n  <w:LsdException Locked=\"false\" Priority=\"64\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Shading 2 Accent 6\"/>\n  <w:LsdException Locked=\"false\" Priority=\"65\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium List 1 Accent 6\"/>\n  <w:LsdException Locked=\"false\" Priority=\"66\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium List 2 Accent 6\"/>\n  <w:LsdException Locked=\"false\" Priority=\"67\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Grid 1 Accent 6\"/>\n  <w:LsdException Locked=\"false\" Priority=\"68\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Grid 2 Accent 6\"/>\n  <w:LsdException Locked=\"false\" Priority=\"69\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Medium Grid 3 Accent 6\"/>\n  <w:LsdException Locked=\"false\" Priority=\"70\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Dark List Accent 6\"/>\n  <w:LsdException Locked=\"false\" Priority=\"71\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Colorful Shading Accent 6\"/>\n  <w:LsdException Locked=\"false\" Priority=\"72\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Colorful List Accent 6\"/>\n  <w:LsdException Locked=\"false\" Priority=\"73\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" Name=\"Colorful Grid Accent 6\"/>\n  <w:LsdException Locked=\"false\" Priority=\"19\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" QFormat=\"true\" Name=\"Subtle Emphasis\"/>\n  <w:LsdException Locked=\"false\" Priority=\"21\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" QFormat=\"true\" Name=\"Intense Emphasis\"/>\n  <w:LsdException Locked=\"false\" Priority=\"31\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" QFormat=\"true\" Name=\"Subtle Reference\"/>\n  <w:LsdException Locked=\"false\" Priority=\"32\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" QFormat=\"true\" Name=\"Intense Reference\"/>\n  <w:LsdException Locked=\"false\" Priority=\"33\" SemiHidden=\"false\"\n   UnhideWhenUsed=\"false\" QFormat=\"true\" Name=\"Book Title\"/>\n  <w:LsdException Locked=\"false\" Priority=\"37\" Name=\"Bibliography\"/>\n  <w:LsdException Locked=\"false\" Priority=\"39\" QFormat=\"true\" Name=\"TOC Heading\"/>\n </w:LatentStyles>\n</xml><![endif]--><!-- [if gte mso 10]>\n<style>\n /* Style Definitions */\n table.MsoNormalTable\n\t{mso-style-name:\"Table Normal\";\n\tmso-tstyle-rowband-size:0;\n\tmso-tstyle-colband-size:0;\n\tmso-style-noshow:yes;\n\tmso-style-priority:99;\n\tmso-style-parent:\"\";\n\tmso-padding-alt:0cm 5.4pt 0cm 5.4pt;\n\tmso-para-margin-top:0cm;\n\tmso-para-margin-right:0cm;\n\tmso-para-margin-bottom:10.0pt;\n\tmso-para-margin-left:0cm;\n\tline-height:115%;\n\tmso-pagination:widow-orphan;\n\tfont-size:11.0pt;\n\tfont-family:\"Calibri\",\"sans-serif\";\n\tmso-ascii-font-family:Calibri;\n\tmso-ascii-theme-font:minor-latin;\n\tmso-hansi-font-family:Calibri;\n\tmso-hansi-theme-font:minor-latin;\n\tmso-fareast-language:EN-US;}\n</style>\n<![endif]--></p>", "gaps": [{"type": "matrix", "useCustomName": true, "customName": "a+b", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "correctAnswer": "{a}+{b}", "correctAnswerFractions": false, "numRows": "2", "numColumns": "2", "allowResize": false, "tolerance": 0, "markPerCell": true, "allowFractions": false, "minColumns": 1, "maxColumns": 0, "minRows": 1, "maxRows": 0}], "sortAnswers": false}, {"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>$\\var{p}\\var{c}-\\var{q}\\var{d}=$[[0]]<!-- [if gte mso 9]><xml>\n <o:OfficeDocumentSettings>\n  <o:AllowPNG/>\n </o:OfficeDocumentSettings>\n</xml><![endif]--></p>\n<p></p>", "gaps": [{"type": "matrix", "useCustomName": true, "customName": "p*c-qd", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "correctAnswer": "{p}*{c}-{q}*{d}", "correctAnswerFractions": false, "numRows": "2", "numColumns": "2", "allowResize": false, "tolerance": 0, "markPerCell": false, "allowFractions": false, "minColumns": 1, "maxColumns": 0, "minRows": 1, "maxRows": 0}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always"}, {"name": "L1 - Matrix Multiplication (2x3 by 3x2 matrices)", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false}, "contributors": [{"name": "Violeta CIT", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/1030/"}, {"name": "Marie Nicholson", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/1799/"}, {"name": "Xiaodan Leng", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2146/"}, {"name": "Timur Zaripov", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/3272/"}], "functions": {}, "rulesets": {"std": ["all", "fractionNumbers", "!collectNumbers"]}, "metadata": {"licence": "Creative Commons Attribution 4.0 International", "description": "<p>Multiplication of&nbsp; matrices of different sizes.</p>"}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["a21", "a22", "b22", "b21", "b12", "b11", "a11", "a12", "a13", "a23", "b31", "b32", "a", "b"], "parts": [{"adaptiveMarkingPenalty": 0, "marks": 0, "customMarkingAlgorithm": "", "showFeedbackIcon": true, "variableReplacements": [], "type": "gapfill", "extendBaseMarkingAlgorithm": true, "unitTests": [], "gaps": [{"adaptiveMarkingPenalty": 0, "marks": "1", "allowResize": false, "customMarkingAlgorithm": "", "showFeedbackIcon": true, "variableReplacements": [], "numRows": "2", "type": "matrix", "correctAnswerFractions": false, "extendBaseMarkingAlgorithm": true, "numColumns": "2", "unitTests": [], "showCorrectAnswer": true, "customName": "ab", "useCustomName": true, "correctAnswer": "{a}*{b}", "scripts": {}, "tolerance": 0, "allowFractions": false, "markPerCell": false, "variableReplacementStrategy": "originalfirst"}], "showCorrectAnswer": true, "customName": "", "sortAnswers": false, "useCustomName": false, "scripts": {}, "prompt": "<p>$AB = \\var{a}\\times\\var{b} = $ [[0]]</p>", "variableReplacementStrategy": "originalfirst"}, {"adaptiveMarkingPenalty": 0, "marks": 0, "customMarkingAlgorithm": "", "showFeedbackIcon": true, "variableReplacements": [], "type": "gapfill", "extendBaseMarkingAlgorithm": true, "unitTests": [], "gaps": [{"adaptiveMarkingPenalty": 0, "marks": "1", "allowResize": false, "customMarkingAlgorithm": "", "showFeedbackIcon": true, "variableReplacements": [], "numRows": "3", "type": "matrix", "correctAnswerFractions": false, "extendBaseMarkingAlgorithm": true, "numColumns": "3", "unitTests": [], "showCorrectAnswer": true, "customName": "ba", "useCustomName": true, "correctAnswer": "{b}*{a}", "scripts": {}, "tolerance": 0, "allowFractions": false, "markPerCell": false, "variableReplacementStrategy": "originalfirst"}], "showCorrectAnswer": true, "customName": "", "sortAnswers": false, "useCustomName": false, "scripts": {}, "prompt": "<p>$BA = \\var{b}\\times\\var{a} = $ [[0]]</p>", "variableReplacementStrategy": "originalfirst"}], "variable_groups": [], "preamble": {"js": "", "css": ""}, "variables": {"b": {"name": "b", "definition": "matrix([b11,b12],[b21,b22],[b31,b32])", "templateType": "anything", "description": "", "group": "Ungrouped variables"}, "a11": {"name": "a11", "definition": "random(-2,1,2)", "templateType": "anything", "description": "", "group": "Ungrouped variables"}, "a23": {"name": "a23", "definition": "random(-2..3)", "templateType": "anything", "description": "", "group": "Ungrouped variables"}, "b12": {"name": "b12", "definition": "random(-3..1)", "templateType": "anything", "description": "", "group": "Ungrouped variables"}, "a": {"name": "a", "definition": "matrix([a11,a12,a13],[a21,a22,a23])", "templateType": "anything", "description": "", "group": "Ungrouped variables"}, "b11": {"name": "b11", "definition": "random(-3,-1,0,3)", "templateType": "anything", "description": "", "group": "Ungrouped variables"}, "b22": {"name": "b22", "definition": "random(-3..-1)", "templateType": "anything", "description": "", "group": "Ungrouped variables"}, "a22": {"name": "a22", "definition": "random(1..3)", "templateType": "anything", "description": "", "group": "Ungrouped variables"}, "a21": {"name": "a21", "definition": "random(-2..2)", "templateType": "anything", "description": "", "group": "Ungrouped variables"}, "b31": {"name": "b31", "definition": "random(-1..1)", "templateType": "anything", "description": "", "group": "Ungrouped variables"}, "b21": {"name": "b21", "definition": "random(2,3)", "templateType": "anything", "description": "", "group": "Ungrouped variables"}, "b32": {"name": "b32", "definition": "random(2..5)", "templateType": "anything", "description": "", "group": "Ungrouped variables"}, "a13": {"name": "a13", "definition": "random(0..2)", "templateType": "anything", "description": "", "group": "Ungrouped variables"}, "a12": {"name": "a12", "definition": "random(1..4)", "templateType": "anything", "description": "", "group": "Ungrouped variables"}}, "tags": [], "advice": "", "statement": "<p>Let <br/>\\[A=\\var{a},\\;\\; B=\\var{b},\\;\\; \\]<br/>Calculate the following products of these matrices:</p>"}, {"name": "Solve a system of three simultaneous linear equations", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false}, "contributors": [{"name": "Frank Doheny", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/789/"}, {"name": "Xiaodan Leng", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2146/"}], "advice": "<p style=\"padding-left: 30px; text-align: left;\">(i) &nbsp; &nbsp;\\(\\var{a1}x+2y+4z=\\var{r1}\\)</p>\n<p style=\"padding-left: 30px;\">(ii) &nbsp; \\(2x+\\var{b1}y+3z=\\var{r2}\\)</p>\n<p style=\"padding-left: 30px;\">(iii) &nbsp;\\(5x+6y+\\var{c1}z=\\var{r3}\\)</p>\n<p>First reduce the three equations in three unknowns to a two equations in two unknowns problem by eliminating one of the variables.</p>\n<p>We can eliminate \\(x\\) using equations (i) and (ii)</p>\n<p style=\"padding-left: 30px;\">2*(i) &nbsp; &nbsp; \\(\\simplify{2*{a1}}x+4y+8z=\\simplify{2*{r1}}\\)</p>\n<p style=\"padding-left: 30px;\">\\(\\var{a1}\\)*(ii) &nbsp; &nbsp;\\(\\simplify{2*{a1}}x+\\simplify{{a1}*{b1}}y+\\simplify{3*{a1}}z=\\simplify{{a1}*{r2}}\\)</p>\n<p>Subtracting gives us a new equation</p>\n<p style=\"padding-left: 30px;\">(iv) &nbsp; &nbsp;\\(\\simplify{(4-{a1}{b1})y+(8-3*{a1})z}=\\simplify{2*{r1}-{a1}*{r2}}\\)</p>\n<p>We can also eliminate \\(x\\) using equations (ii) and (iii)</p>\n<p style=\"padding-left: 30px;\">5*(ii) &nbsp; &nbsp;\\(10x +\\simplify{5*{b1}}y+15z=\\simplify{5*{r2}}\\)</p>\n<p style=\"padding-left: 30px;\">2*(iii) &nbsp; \\(10x+12y+\\simplify{2*{c1}}z=\\simplify{2*{r3}}\\)</p>\n<p>Subtracting gives us another new equation</p>\n<p style=\"padding-left: 30px;\">(v) &nbsp; &nbsp; \\(\\simplify{(5*{b1}-12)y+(15-2*{c1})z}=\\simplify{5*{r2}-2*{r3}}\\)</p>\n<p>We could then eliminate the \\(y\\) from these two new equations</p>\n<p style=\"padding-left: 30px;\">\\(\\simplify{5*{b1}-12}\\)*(iv) &nbsp; &nbsp;\\(\\simplify{(5*{b1}-12)*(4-{a1}{b1})y+(5*{b1}-12)*(8-3*{a1})z}=\\simplify{(5*{b1}-12)*(2*{r1}-{a1}*{r2})}\\)</p>\n<p style=\"padding-left: 30px;\">\\(\\simplify{4-{a1}{b1}}\\)*(v) &nbsp; &nbsp;\\(\\simplify{(4-{a1}{b1})*(5*{b1}-12)y+(4-{a1}{b1})*(15-2*{c1})z}=\\simplify{(4-{a1}{b1})*(5*{r2}-2*{r3})}\\)</p>\n<p>Subtracting gives us</p>\n<p style=\"padding-left: 30px;\">\\(\\simplify{(5*{b1}-12)*(8-3*{a1})-(4-{a1}{b1})*(15-2*{c1})}z=\\simplify{(5*{b1}-12)*(2*{r1}-{a1}*{r2})-(4-{a1}{b1})*(5*{r2}-2*{r3})}\\)</p>\n<p>Thus</p>\n<p style=\"padding-left: 30px;\">\\(z=\\frac{\\simplify{(5*{b1}-12)*(2*{r1}-{a1}*{r2})-(4-{a1}{b1})*(5*{r2}-2*{r3})}}{\\simplify{(5*{b1}-12)*(8-3*{a1})-(4-{a1}{b1})*(15-2*{c1})}}=\\simplify{decimal{((5*{b1}-12)*(2*{r1}-{a1}*{r2})-(4-{a1}*{b1})*(5*{r2}-2*{r3}))/( &nbsp;(5*{b1}-12)*(8-3*{a1})-(4-{a1}*{b1})*(15-2*{c1}))}}\\)</p>\n<p>We can now back substitute this value for \\(z\\) into equation (iv) to find the correct value for \\(y\\) and then&nbsp;back substitute both these values into equation (i) to calculate \\(x\\).&nbsp;</p>\n<p style=\"padding-left: 30px;\"></p>", "variable_groups": [], "ungrouped_variables": ["a1", "b1", "c1", "r1", "r2", "r3"], "functions": {}, "variablesTest": {"condition": "", "maxRuns": 100}, "type": "question", "statement": "<p>Solve the following system of&nbsp;three simultaneous linear equations:</p>\n<p style=\"padding-left: 30px;\">\\(\\var{a1}x+2y+4z=\\var{r1}\\)</p>\n<p style=\"padding-left: 60px;\">and</p>\n<p style=\"padding-left: 30px;\">\\(2x+\\var{b1}y+3z=\\var{r2}\\)</p>\n<p style=\"padding-left: 60px;\">and</p>\n<p style=\"padding-left: 30px;\">\\(5x+6y+\\var{c1}z=\\var{r3}\\)</p>", "preamble": {"js": "", "css": ""}, "tags": [], "parts": [{"showFeedbackIcon": true, "marks": 0, "scripts": {}, "prompt": "<p>Input the value of \\(x\\) that satisfies the three equations.</p>\n<p style=\"padding-left: 30px;\">\\(x = \\)&nbsp;[[0]]</p>\n<p>Input the value of \\(y\\) that satisfies the three equations.</p>\n<p style=\"padding-left: 30px;\">\\(y = \\)&nbsp;[[1]]</p>\n<p>Input the value of \\(z\\) that satisfies the three equations.</p>\n<p style=\"padding-left: 30px;\">\\(z = \\)&nbsp;[[2]]</p>", "type": "gapfill", "gaps": [{"correctAnswerStyle": "plain", "scripts": {}, "strictPrecision": false, "maxValue": "((18-{b1}*{c1})*(3*{r1}-{r3})+({c1}-12)*(6*{r2}-{b1}*{r3}))/((3*{a1}-5)*(18-{b1}*{c1})-(12-5*{b1})*(12-{c1}))", "showFeedbackIcon": true, "showPrecisionHint": true, "precision": "3", "mustBeReducedPC": 0, "showCorrectAnswer": true, "variableReplacementStrategy": "originalfirst", "type": "numberentry", "precisionMessage": "You have not given your answer to the correct precision.", "notationStyles": ["plain", "en", "si-en"], "marks": 1, "precisionType": "dp", "allowFractions": false, "minValue": "((18-{b1}*{c1})*(3*{r1}-{r3})+({c1}-12)*(6*{r2}-{b1}*{r3}))/((3*{a1}-5)*(18-{b1}*{c1})-(12-5*{b1})*(12-{c1}))", "precisionPartialCredit": 0, "correctAnswerFraction": false, "variableReplacements": [], "mustBeReduced": false}, {"correctAnswerStyle": "plain", "scripts": {}, "strictPrecision": false, "maxValue": "((15-2*{c1})*(2*{r1}-{a1}*{r2})+(3*{a1}-8)*(5*{r2}-2*{r3}))/((4-{a1}*{b1})*(15-2*{c1})-(5*{b1}-12)*(8-3*{a1}))", "showFeedbackIcon": true, "showPrecisionHint": true, "precision": "3", "mustBeReducedPC": 0, "showCorrectAnswer": true, "variableReplacementStrategy": "originalfirst", "type": "numberentry", "precisionMessage": "You have not given your answer to the correct precision.", "notationStyles": ["plain", "en", "si-en"], "marks": 1, "precisionType": "dp", "allowFractions": false, "minValue": "((15-2*{c1})*(2*{r1}-{a1}*{r2})+(3*{a1}-8)*(5*{r2}-2*{r3}))/((4-{a1}*{b1})*(15-2*{c1})-(5*{b1}-12)*(8-3*{a1}))", "precisionPartialCredit": 0, "correctAnswerFraction": false, "variableReplacements": [], "mustBeReduced": false}, {"correctAnswerStyle": "plain", "scripts": {}, "strictPrecision": false, "maxValue": "((12-5*{b1})*(2*{r1}-{a1}*{r2})+(4-{b1}*{a1})*(5*{r2}-2*{r3}))/((4-{a1}*{b1})*(15-2*{c1})-(5*{b1}-12)*(8-3*{a1}))", "showFeedbackIcon": true, "showPrecisionHint": true, "precision": "3", "mustBeReducedPC": 0, "showCorrectAnswer": true, "variableReplacementStrategy": "originalfirst", "type": "numberentry", "precisionMessage": "You have not given your answer to the correct precision.", "notationStyles": ["plain", "en", "si-en"], "marks": 1, "precisionType": "dp", "allowFractions": false, "minValue": "((12-5*{b1})*(2*{r1}-{a1}*{r2})+(4-{b1}*{a1})*(5*{r2}-2*{r3}))/((4-{a1}*{b1})*(15-2*{c1})-(5*{b1}-12)*(8-3*{a1}))", "precisionPartialCredit": 0, "correctAnswerFraction": false, "variableReplacements": [], "mustBeReduced": false}], "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "showCorrectAnswer": true}], "rulesets": {}, "variables": {"r3": {"templateType": "randrange", "description": "", "definition": "random(30..60#1)", "group": "Ungrouped variables", "name": "r3"}, "r2": {"templateType": "randrange", "description": "", "definition": "random(18..50#1)", "group": "Ungrouped variables", "name": "r2"}, "r1": {"templateType": "randrange", "description": "", "definition": "random(20..42#1)", "group": "Ungrouped variables", "name": "r1"}, "b1": {"templateType": "randrange", "description": "", "definition": "random(2..10#1)", "group": "Ungrouped variables", "name": "b1"}, "c1": {"templateType": "randrange", "description": "", "definition": "random(3..12#1)", "group": "Ungrouped variables", "name": "c1"}, "a1": {"templateType": "randrange", "description": "", "definition": "random(2..8#1)", "group": "Ungrouped variables", "name": "a1"}}, "metadata": {"description": "<p>Solve a system of three simultaneous linear equations</p>", "licence": "Creative Commons Attribution 4.0 International"}}]}], "allowPrinting": true, "navigation": {"allowregen": true, "reverse": true, "browse": true, "allowsteps": true, "showfrontpage": true, "showresultspage": "oncompletion", "navigatemode": "sequence", "onleave": {"action": "none", "message": ""}, "preventleave": true, "startpassword": ""}, "timing": {"allowPause": true, "timeout": {"action": "none", "message": ""}, "timedwarning": {"action": "none", "message": ""}}, "feedback": {"showactualmark": true, "showtotalmark": true, "showanswerstate": true, "allowrevealanswer": true, "advicethreshold": 0, "intro": "", "reviewshowscore": true, "reviewshowfeedback": true, "reviewshowexpectedanswer": true, "reviewshowadvice": true, "feedbackmessages": []}, "diagnostic": {"knowledge_graph": {"topics": [], "learning_objectives": []}, "script": "diagnosys", "customScript": ""}, "contributors": [{"name": "abc temp", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/18152/"}], "extensions": ["stats"], "custom_part_types": [], "resources": []}