// Numbas version: exam_results_page_options
{"name": "Jesse's copy of Numbas demo: match choices with answers part type", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Jesse's copy of Numbas demo: match choices with answers part type", "tags": [], "metadata": {"description": "<p>A demo of the match choices with answers part and its options.</p>", "licence": "Creative Commons Attribution 4.0 International"}, "statement": "<p>The <em>match choices with answers</em> part presents the student with a list of <em>choices</em>, and a list of <em>answers</em>. The student must match up choices with answers, either choosing one answer for each choice, or arbitrarily many pairs.</p>\n<p>Each choice-answer pair awards a different number of marks.</p>\n<p><em><a href=\"https://numbas.mathcentre.ac.uk/question/66215/numbas-demo-match-choices-with-answers-part-type/\" target=\"_blank\" rel=\"noopener noreferrer\">See this question on the public editor</a>.</em></p>", "advice": "", "rulesets": {}, "extensions": [], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true}, "constants": [], "variables": {"target": {"name": "target", "group": "Ungrouped variables", "definition": "random(2,3,4,4,5,6)", "description": "", "templateType": "anything", "can_override": false}, "random_marking_matrix": {"name": "random_marking_matrix", "group": "Ungrouped variables", "definition": "map(\n  map(award(1,a+b=target),b,1..3),\n  a,\n  1..3\n)", "description": "", "templateType": "anything", "can_override": false}, "number_of_choices": {"name": "number_of_choices", "group": "Ungrouped variables", "definition": "random(2..4)", "description": "", "templateType": "anything", "can_override": false}, "random_choices": {"name": "random_choices", "group": "Ungrouped variables", "definition": "offset+1..offset+number_of_choices", "description": "", "templateType": "anything", "can_override": false}, "offset": {"name": "offset", "group": "Ungrouped variables", "definition": "random(0..8)", "description": "", "templateType": "anything", "can_override": false}, "even_odd_marking_matrix": {"name": "even_odd_marking_matrix", "group": "Ungrouped variables", "definition": "map(\n  if(mod(j,2)=0,[1,0],[0,1]),\n  j,\n  random_choices\n)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["target", "random_marking_matrix", "number_of_choices", "offset", "random_choices", "even_odd_marking_matrix"], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "m_n_x", "useCustomName": true, "customName": "One from each row", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "minAnswers": 0, "maxAnswers": 0, "shuffleChoices": false, "shuffleAnswers": false, "displayType": "radiogroup", "warningType": "none", "showCellAnswerState": true, "markingMethod": "sum ticked cells", "choices": ["Choice 1", "Choice 2"], "matrix": [["1", 0], [0, "1"]], "layout": {"type": "all", "expression": ""}, "answers": ["Answer 1", "Answer 2"]}, {"type": "m_n_x", "useCustomName": true, "customName": "Checkboxes", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "minAnswers": 0, "maxAnswers": 0, "shuffleChoices": false, "shuffleAnswers": false, "displayType": "checkbox", "warningType": "none", "showCellAnswerState": true, "choices": ["Choice A", "Choice B"], "matrix": [["1", "1", 0, 0], [0, 0, "1", "1"]], "layout": {"type": "all", "expression": ""}, "answers": ["Answer A1", "Answer A2", "Answer B1", "Answer B2"]}, {"type": "m_n_x", "useCustomName": true, "customName": "Shuffle the order of choices and answers", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>You can randomise the order of both choices and answers, or just one.</p>", "minMarks": 0, "maxMarks": 0, "minAnswers": 0, "maxAnswers": 0, "shuffleChoices": true, "shuffleAnswers": true, "displayType": "checkbox", "warningType": "none", "showCellAnswerState": true, "choices": ["Choice A", "Choice B"], "matrix": [["1", "1", 0, 0], [0, 0, "1", "1"]], "layout": {"type": "all", "expression": ""}, "answers": ["Answer A1", "Answer A2", "Answer B1", "Answer B2"]}, {"type": "m_n_x", "useCustomName": true, "customName": "Remove some options", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p><em>When certain selections don't make sense, or there's symmetry, you can remove them by choosing a different layout.</em></p>\n<p>Select pairs of numbers that add up to 4.</p>", "minMarks": 0, "maxMarks": 0, "minAnswers": 0, "maxAnswers": 0, "shuffleChoices": false, "shuffleAnswers": false, "displayType": "checkbox", "warningType": "none", "showCellAnswerState": true, "choices": ["1", "2", "3"], "matrix": [[0, 0, 0], ["0", "3", 0], ["3", 0, 0]], "layout": {"type": "lowertriangle", "expression": ""}, "answers": ["1", "2", "3"]}, {"type": "m_n_x", "useCustomName": true, "customName": "Marking depends on question variables", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p><em>The marks awarded for each selection can depend on question variables.</em></p>\n<p>This time, select pairs of numbers that add up to <strong>{target}</strong>.</p>", "minMarks": 0, "maxMarks": 0, "minAnswers": 0, "maxAnswers": 0, "shuffleChoices": false, "shuffleAnswers": false, "displayType": "checkbox", "warningType": "none", "showCellAnswerState": true, "markingMethod": "sum ticked cells", "choices": ["1", "2", "3"], "matrix": "random_marking_matrix", "layout": {"type": "lowertriangle", "expression": ""}, "answers": ["1", "2", "3"]}, {"type": "m_n_x", "useCustomName": true, "customName": "Variable number of choices", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p><em>Even the lists of choices and answers can depend on question variables.</em></p>\n<p>Classify each of the following numbers as even or odd.</p>", "minMarks": 0, "maxMarks": 0, "minAnswers": 0, "maxAnswers": 0, "shuffleChoices": false, "shuffleAnswers": false, "displayType": "radiogroup", "warningType": "none", "showCellAnswerState": true, "markingMethod": "sum ticked cells", "choices": "map(string(n),n,random_choices)", "matrix": "even_odd_marking_matrix", "layout": {"type": "all", "expression": ""}, "answers": ["Even", "Odd"]}, {"type": "m_n_x", "useCustomName": true, "customName": "Score per matched cell", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p><em>This part uses the \"score per matched cell\" marking method - you will be awarded 1 mark for each correct ticked or unticked pair.</em></p>\n<p>Tick each pair of coprime numbers. Two numbers are coprime if their greatest common factor is 1.</p>", "minMarks": 0, "maxMarks": "12", "minAnswers": 0, "maxAnswers": 0, "shuffleChoices": true, "shuffleAnswers": false, "displayType": "checkbox", "warningType": "none", "showCellAnswerState": true, "markingMethod": "score per matched cell", "choices": ["6", "4", "14"], "matrix": [[0, 0, "1", "1"], [0, "1", "1", "1"], [0, "1", 0, 0]], "layout": {"type": "all", "expression": ""}, "answers": ["18", "9", "7", "35"]}, {"type": "m_n_x", "useCustomName": true, "customName": "All-or-nothing score", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p><em>This part uses the \"all-or-nothing\" marking method - you will be awarded the maxMarks if you tick each box correctly</em></p>\n<p>Tick each pair of coprime numbers. Two numbers are coprime if their greatest common factor is 1.</p>", "minMarks": 0, "maxMarks": "12", "minAnswers": 0, "maxAnswers": 0, "shuffleChoices": true, "shuffleAnswers": false, "displayType": "checkbox", "warningType": "none", "showCellAnswerState": true, "markingMethod": "all-or-nothing", "choices": ["6", "4", "14"], "matrix": [[0, 0, "1", "1"], [0, "1", "1", "1"], [0, "1", 0, 0]], "layout": {"type": "all", "expression": ""}, "answers": ["18", "9", "7", "35"]}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "Christian Lawson-Perfect", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/7/"}, {"name": "Jesse Hoobergs", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/12130/"}]}]}], "contributors": [{"name": "Christian Lawson-Perfect", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/7/"}, {"name": "Jesse Hoobergs", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/12130/"}]}