// Numbas version: exam_results_page_options
{"name": "Jesse's copy of Numbas demo: choose several from a list part type", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Jesse's copy of Numbas demo: choose several from a list part type", "tags": [], "metadata": {"description": "<p>A demo of the choose several from a list part and its options.</p>", "licence": "Creative Commons Attribution 4.0 International"}, "statement": "<p>The <em>choose several from a list</em> part type presents the student with a list of choices, of which they have to select one or more. Each choice awards a different number of marks when it is chosen, and the student is awarded the sum of the marks for their choices.</p>\n<p><em><a href=\"https://numbas.mathcentre.ac.uk/question/66212/numbas-demo-choose-several-from-a-list-part-type/\" target=\"_blank\" rel=\"noopener noreferrer\">See this question on the public editor</a>.</em></p>", "advice": "", "rulesets": {}, "extensions": [], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true}, "constants": [], "variables": {"pick_odd_numbers": {"name": "pick_odd_numbers", "group": "Ungrouped variables", "definition": "random(true,false)", "description": "", "templateType": "anything", "can_override": false}, "numbers_to_pick": {"name": "numbers_to_pick", "group": "Ungrouped variables", "definition": "if(pick_odd_numbers,\"odd\",\"even\")", "description": "", "templateType": "anything", "can_override": false}, "numbers_marking_matrix": {"name": "numbers_marking_matrix", "group": "Ungrouped variables", "definition": "map(\n  award(1,mod(j,2) = if(pick_odd_numbers,1,0)),\n  j,\n  1..4\n)", "description": "", "templateType": "anything", "can_override": false}, "number_of_correct_choices": {"name": "number_of_correct_choices", "group": "Ungrouped variables", "definition": "random(2..total_choices-2 except 5)", "description": "", "templateType": "anything", "can_override": false}, "number_of_distractors": {"name": "number_of_distractors", "group": "Ungrouped variables", "definition": "total_choices - number_of_correct_choices", "description": "", "templateType": "anything", "can_override": false}, "random_choices": {"name": "random_choices", "group": "Ungrouped variables", "definition": "map(\"Correct choice \"+j,j,1..number_of_correct_choices)\n+\nmap(\"Incorrect choice \"+j,j,1..number_of_distractors)", "description": "", "templateType": "anything", "can_override": false}, "random_marking_matrix": {"name": "random_marking_matrix", "group": "Ungrouped variables", "definition": "repeat(12/number_of_correct_choices,number_of_correct_choices) + repeat(-12/number_of_correct_choices,number_of_distractors)", "description": "", "templateType": "anything", "can_override": false}, "total_choices": {"name": "total_choices", "group": "Ungrouped variables", "definition": "random(4..7)", "description": "", "templateType": "anything", "can_override": false}, "one": {"name": "one", "group": "Ungrouped variables", "definition": "0", "description": "", "templateType": "number", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["pick_odd_numbers", "numbers_to_pick", "numbers_marking_matrix", "total_choices", "number_of_correct_choices", "number_of_distractors", "random_choices", "random_marking_matrix", "one"], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "m_n_2", "useCustomName": true, "customName": "Award marks for correct choices", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>In this part, the student is awarded 1 mark for each correct choice. However, selecting all the choices also gives you full marks...</p>", "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "checkbox", "displayColumns": 0, "minAnswers": 0, "maxAnswers": 0, "warningType": "none", "showCellAnswerState": true, "markingMethod": "sum ticked cells", "choices": ["Correct choice", "Correct choice", "Incorrect choice", "Incorrect choice"], "matrix": ["1", "1", 0, 0], "distractors": ["", "", "", ""]}, {"type": "m_n_2", "useCustomName": true, "customName": "Display choices in columns", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>You can display the choices in columns.</p>", "minMarks": "0", "maxMarks": "0", "shuffleChoices": false, "displayType": "checkbox", "displayColumns": "3", "minAnswers": "0", "maxAnswers": "0", "warningType": "prevent", "showCellAnswerState": true, "markingMethod": "sum ticked cells", "choices": ["Correct choice", "Correct choice", "Incorrect choice", "Incorrect choice", "Incorrect choice", "Incorrect choice"], "matrix": ["1", "1", 0, 0, 0, 0], "distractors": ["", "", "", "", "", ""]}, {"type": "m_n_2", "useCustomName": true, "customName": "Restricted number of answers", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>You must pick exactly two answers.</p>", "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "checkbox", "displayColumns": "1", "minAnswers": "2", "maxAnswers": "2", "warningType": "prevent", "showCellAnswerState": true, "markingMethod": "sum ticked cells", "choices": ["Correct choice", "Correct choice", "Incorrect choice", "Incorrect choice"], "matrix": ["1", "1", 0, 0], "distractors": ["", "", "", ""]}, {"type": "m_n_2", "useCustomName": true, "customName": "Negative marking", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>You can take away marks when incorrect choices are selected, to discourage the student from picking at random.</p>\n<p>Each incorrect choice for this part takes away 2 marks.</p>", "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "checkbox", "displayColumns": "1", "minAnswers": 0, "maxAnswers": 0, "warningType": "none", "showCellAnswerState": true, "markingMethod": "sum ticked cells", "choices": ["Correct choice", "Correct choice", "Correct choice", "Incorrect choice", "Incorrect choice"], "matrix": ["1", "1", "1", "-2", "-2"], "distractors": ["", "", "", "", ""]}, {"type": "m_n_2", "useCustomName": true, "customName": "Score per matched cell", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p><em>This part uses the \"score per matched cell\" marking method - you will be awarded 1 mark for each correct ticked or unticked box.</em></p>\n<p>Tick each mammal.</p>", "minMarks": 0, "maxMarks": "2", "shuffleChoices": true, "displayType": "checkbox", "displayColumns": 0, "minAnswers": 0, "maxAnswers": 0, "warningType": "none", "showCellAnswerState": true, "markingMethod": "score per matched cell", "choices": ["Cat", "Dog", "Shark", "Sparrow"], "matrix": ["1", "1", 0, 0], "distractors": ["", "", "", ""]}, {"type": "m_n_2", "useCustomName": true, "customName": "All-or-nothing score", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p><em>This part uses the \"all-or-nothing\" marking method - you will be awarded the Maximum marks if you tick exactly the right boxes.</em></p>\n<p>Tick each mammal.</p>", "minMarks": 0, "maxMarks": "5", "shuffleChoices": true, "displayType": "checkbox", "displayColumns": 0, "minAnswers": 0, "maxAnswers": 0, "warningType": "none", "showCellAnswerState": true, "markingMethod": "all-or-nothing", "choices": ["Cat", "Dog", "Shark", "Sparrow"], "matrix": ["1", "1", 0, 0], "distractors": ["", "", "", ""]}, {"type": "m_n_2", "useCustomName": true, "customName": "Distractor text", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>For each choice, you can write a message to display to the student if they choose it.</p>", "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "checkbox", "displayColumns": "1", "minAnswers": 0, "maxAnswers": 0, "warningType": "none", "showCellAnswerState": true, "choices": ["Correct choice 1", "Correct choice 1", "Incorrect choice 1", "Incorrect choice 2"], "matrix": ["1", "1", "-1", "-1"], "distractors": ["Correct choice 1 is indeed correct!", "Correct choice 2 is indeed correct!", "Incorrect choice 1 is not correct.", "Incorrect choice 2 is not correct."]}, {"type": "m_n_2", "useCustomName": true, "customName": "Shuffle the order of choices", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>You can randomly shuffle the order of the choices.</p>\n<p><em>Click the \"Try another question like this one\" button to see these choices in a different order.</em></p>", "minMarks": 0, "maxMarks": 0, "shuffleChoices": true, "displayType": "checkbox", "displayColumns": "1", "minAnswers": 0, "maxAnswers": 0, "warningType": "none", "showCellAnswerState": true, "markingMethod": "sum ticked cells", "choices": ["Correct choice A", "Correct choice B", "Correct choice C", "Incorrect choice 1", "Incorrect choice 2"], "matrix": ["1", "1", "1", 0, 0], "distractors": ["", "", "", "", ""]}, {"type": "m_n_2", "useCustomName": true, "customName": "Marking depends on variables", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>You can set the number of marks based on question variables.</p>\n<p>This time, pick <strong>{numbers_to_pick}</strong> numbers.</p>", "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "checkbox", "displayColumns": 0, "minAnswers": 0, "maxAnswers": 0, "warningType": "none", "showCellAnswerState": true, "markingMethod": "sum ticked cells", "choices": ["1", "2", "3", "4"], "matrix": "numbers_marking_matrix"}, {"type": "m_n_2", "useCustomName": true, "customName": "Variable number of choices", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p><em>The choices themselves can depend on question variables, too.</em></p>\n<p>This time, there are {number_of_correct_choices} correct choices and {number_of_distractors} incorrect choices.</p>", "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "checkbox", "displayColumns": "1", "minAnswers": 0, "maxAnswers": 0, "warningType": "none", "showCellAnswerState": true, "markingMethod": "sum ticked cells", "choices": "random_choices", "matrix": "random_marking_matrix"}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "Christian Lawson-Perfect", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/7/"}, {"name": "Jesse Hoobergs", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/12130/"}]}]}], "contributors": [{"name": "Christian Lawson-Perfect", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/7/"}, {"name": "Jesse Hoobergs", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/12130/"}]}