// Numbas version: exam_results_page_options
{"name":"First Exam","duration":0,"percentPass":50.0,"showQuestionGroupNames":true,"showstudentname":false,"allowPrinting":false,"resources":[],"extensions":[],"custom_part_types":[],"navigation":{"startpassword":"","allowregen":true,"navigatemode":"sequence","onleave":{"action":"none","message":""},"showresultspage":"never","reverse":true,"browse":false,"allowsteps":false,"showfrontpage":false,"preventleave":false},"timing":{"allowPause":true,"timeout":{"action":"none","message":""},"timedwarning":{"action":"none","message":""}},"feedback":{"showactualmark":false,"showtotalmark":true,"showanswerstate":true,"allowrevealanswer":true,"reviewshowscore":true,"reviewshowfeedback":true,"reviewshowexpectedanswer":true,"reviewshowadvice":true,"advice":"Please practice more","intro":"An introductory text","feedbackmessages":[]},"question_groups":[{"name":"Group 1","pickingStrategy":"all-ordered","questions":[{"name":"first_question.yaml","statement":"How much is 9 * 3?","advice":"You could calculate 10 * 3 and then subtract 3. This yield 3 * 10 - 3 = 30 - 3 = 27","parts":[{"type":"numberentry","prompt":"Enter your answer in the inputbox","marks":5,"steps":[],"stepsPenalty":0,"showCorrectAnswer":true,"showFeedbackIcon":true,"useCustomName":false,"customName":"","variableReplacements":[],"variableReplacementStrategy":"originalfirst","adaptiveMarkingPenalty":0,"customMarkingAlgorithm":"","extendBaseMarkingAlgorithm":true,"correctAnswerFraction":false,"correctAnswerStyle":"plain","allowFractions":false,"notationStyles":["plain"],"mustBeReduced":false,"mustBeReducedPC":0,"precisionType":"none","precision":0,"precisionPartialCredit":0,"precisionMessage":"","strictPrecision":true,"showPrecisionHint":true,"showFractionHint":false,"answer":"27"}],"builtin_constants":{"e":false,"i":false,"pi,π":false},"constants":[],"variables":{},"variablesTest":{"condition":"","maxRuns":1},"functions":{},"ungrouped_variables":[],"variable_groups":[],"rulesets":{},"preamble":{"js":"","css":""},"navigation":{"allowregen":false,"showfrontpage":false,"preventleave":false},"extensions":[],"resources":[],"tags":[],"custom_part_types":[]},{"name":"first_question_with_variables.yaml","statement":"How much is 9 * {a}?","advice":"You could calculate 10 * {a} and then subtract {a}. This yield {a} * 10 - {a} = {a * 10} - {a} = {a*9}","parts":[{"type":"numberentry","prompt":"Enter your answer in the inputbox","marks":5,"steps":[],"stepsPenalty":0,"showCorrectAnswer":true,"showFeedbackIcon":true,"useCustomName":false,"customName":"","variableReplacements":[],"variableReplacementStrategy":"originalfirst","adaptiveMarkingPenalty":0,"customMarkingAlgorithm":"","extendBaseMarkingAlgorithm":true,"correctAnswerFraction":false,"correctAnswerStyle":"plain","allowFractions":false,"notationStyles":["plain"],"mustBeReduced":false,"mustBeReducedPC":0,"precisionType":"none","precision":0,"precisionPartialCredit":0,"precisionMessage":"","strictPrecision":true,"showPrecisionHint":true,"showFractionHint":false,"answer":"a * 9"}],"builtin_constants":{"e":false,"i":false,"pi,π":false},"constants":[],"variables":{"a":{"name":"a","definition":"3","description":"","templateType":"number","group":"Ungrouped variables","can_override":false}},"variablesTest":{"condition":"","maxRuns":1},"functions":{},"ungrouped_variables":["a"],"variable_groups":[],"rulesets":{},"preamble":{"js":"","css":""},"navigation":{"allowregen":false,"showfrontpage":false,"preventleave":false},"extensions":[],"resources":[],"tags":[],"custom_part_types":[]},{"name":"first_question_with_randomisation.yaml","statement":"How much is 9 * {a}?","advice":"You could calculate 10 * {a} and then subtract {a}. This yield {a} * 10 - {a} = {a * 10} - {a} = {a*9}","parts":[{"type":"numberentry","prompt":"Enter your answer in the inputbox","marks":5,"steps":[],"stepsPenalty":0,"showCorrectAnswer":true,"showFeedbackIcon":true,"useCustomName":false,"customName":"","variableReplacements":[],"variableReplacementStrategy":"originalfirst","adaptiveMarkingPenalty":0,"customMarkingAlgorithm":"","extendBaseMarkingAlgorithm":true,"correctAnswerFraction":false,"correctAnswerStyle":"plain","allowFractions":false,"notationStyles":["plain"],"mustBeReduced":false,"mustBeReducedPC":0,"precisionType":"none","precision":0,"precisionPartialCredit":0,"precisionMessage":"","strictPrecision":true,"showPrecisionHint":true,"showFractionHint":false,"answer":"a * 9"}],"builtin_constants":{"e":false,"i":false,"pi,π":false},"constants":[],"variables":{"a":{"name":"a","definition":"random(1 .. 10#1)","description":"","templateType":"randrange","group":"Ungrouped variables","can_override":false}},"variablesTest":{"condition":"","maxRuns":1},"functions":{},"ungrouped_variables":["a"],"variable_groups":[],"rulesets":{},"preamble":{"js":"","css":""},"navigation":{"allowregen":false,"showfrontpage":false,"preventleave":false},"extensions":[],"resources":[],"tags":[],"custom_part_types":[]}]}]}